Why the brown ghost chirps at night

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

The authors investigate the role of chirping in a species of weakly electric fish. They subject the fish to various scenarios and correlate the production of chirps with many different factors. They find major correlations between the background beat signals (continuously present during any social interactions) or some aspects of social and environmental conditions with the propensity to produce different types of chirps. By analyzing more specifically different aspects of these correlations they conclude that chirping patterns are related to navigation purposes and the need to localize the source of the beat signal (i.e. the location of the conspecific).

The study provides a wealth of interesting observations of behavior and much of this data constitutes a useful dataset to document the patterns of social interactions in these fish. Some data, in particular the high propensity to chirp in cluttered environments, raises interesting questions. Their main hypothesis is a useful addition to the debate on the function of these chirps and is worth considering and exploring further.

After the initial reviewers' comments, the authors performed a welcome revision of the way the results are presented. Overall the study has been improved by the revision. However, one piece of new data is perplexing to me. The new Figure 7 presents the results of a model analysis of the strength of the EI caused by a second fish to localize when the focal fish is chirping. From my understanding of this type of model, EOD frequency is not a parameter in the model since it evaluates the strength of the field at a given point in time. Therefore the only thing that matters is the phase relationship and strength of the EOD. Assuming that the second fish's EOD is kept constant and the phases relationship is also the same, the only difference during a chirp that could affect the result of the calculation is the potential decrease in EOD amplitude during the chirp. It is indeed logical that if the focal fish decreased its EOD amplitude the target fish's EOD becomes relatively stronger. Where things are harder to understand is why the different types of chirps (e.g. type 1 vs type 2) lead to the same increase in signal even though they are typically associated with different levels of amplitude modulations. Also, it is hard to imagine that a type 2 chirps that is barely associated with any decrease in EOD amplitude (0-10% maybe), would cause doubling of the EI strength. There might be something I don't understand but the authors should provide a lot more details on how this result is obtained and convince us that it makes sense.

We thank the author for the comments and we agree that the approach could have been better detailed. As anticipated by the Reviewer, the Boundary Element Method (BEM) model can be used simply to calculate the electric field and electric image at a specific point in time (instantaneously), regardless of EOD frequency. However, our model allows for the concatenation of consecutive instants and thus is able to render an entire sequence of electric fields - and resulting electric images - incorporating realistic EOD characteristics such as shape, duration, and frequencies (see Pedraja et al., 2014).

Chirp-triggered EIs were modeled using real chirps produced by interacting fish. Each chirp was thus associated to its duration and peak parameters, as well as the fish positional information (distance and angle).

However, since we did not know the beat phase at which chirps were produced, we computed electric images for each fish position and chirp scenario by simulating various phases (here referred to the initial offset of the two EODs, set at 4 phases, equally spaced). These are intended as phases of the sender EOD and simply refer to the initial OFFSET between the two interacting EODs. However, since our simulations were run over a time window of 500 msec, all phases are likely to be covered, with a different temporal order relative to the chirp (always centered within the 500 msec).

The simulation was run maintaining consistent timing for both chirp and non-chirp conditions, across approximately 800 body nodes. At each node, the current flow was calculated from the peak-to-peak of the EOD sum (i.e. the point-to-point of the difference between the beat positive and negative envelopes). Analyzing the EIs over this fixed time window enables us to assess the unitary changes of current flow induced by chirps over units of time (ΔI/Δt). From this, we can calculate a cumulative sum of current flow changes - expressed as delta(EI) and use it to show the effect of the chirps on the spatiotemporal EI (Figure 7C).

One can express this cumulative change mapped onto the fish body (keeping the 800 points separated, as in Figure 7C) or further sum the current changes to obtain a single total (as shown in Figure 7D).

One can check this by considering that a sum for example of a set of 500/800 points - judging from the size of the blue areas in C not all 800 points have a detectable change - each valued 0.1-to-0.3 mA/s, one could get circa 100 mA/s, which is what is shown in D. (is this what is happening ?)

We do not know why chirps of different types triggered similar effects. It is possible that, since EI measurements are pooled over several chirps produced at different angles and distances, in case of a lower amount of chirps considered for a given type (as in the case of rises, very low) these measurements may not highlight more marked differences among types. In a publication we are currently working on, we are considering a larger dataset to better assess these results.

The methods section has been edited to clarify the approach (not yet).

Reviewer #2 (Public Review):

Studying Apteronotus leptorhynchus (the weakly electric brown ghost knifefish), the authors provide evidence that 'chirps' (brief modulations in the frequency and amplitude of the ongoing electric signal) function in active sensing (specifically homeoactive sensing) rather than communication. Chirping is a behavior that has been well studied, including numerous studies on the sensory coding of chirps and the neural mechanisms for chirp generation.

Chirps are largely thought to function in communication behavior, so this alternative function is a very exciting possibility that could have a great impact on the field.

We thank the Reviewer for the extensive and constructive comments. We would like to add that, while it is true that many detailed studies have been published on the anatomy and physiology of the circuits implicated in the production and modulation of “electric chirps”, most of this research assumed, and focused exclusively on, their possible role in communication. In addition, most behavioral studies did the same and a meta-analysis of the existing literature on chirping allows to trace back the communication idea mainly to two studies: Hagedorn and Heiligenberg, 1985 (“Court and spark: electric signals in the courtship and mating of gymnotoid fish”) and Hopkins, 1974 (“Electric Communication: Functions in the Social Behavior of Eigenmannia Virescens”), among the main sources. Importantly, in these studies only contextual observations have been made (no playback experiment or other attempts to analyze more quantitatively the correlation of chirping with other behaviors).

The authors do provide convincing evidence that chirps may function in homeoactive sensing. However, their evidence arguing against a role for chirps in communication is not as strong, and fails to sufficiently consider the evidence from a large body of existing research. Ultimately, the manuscript presents very interesting data that is sure to stimulate discussion and follow-up studies, but it suffers from dismissing evidence in support of, or consistent with, a communicative function for chirps.

Although the tone of some statements present in our earlier draft may suggest otherwise, through our revisions, we have made an effort to clarify that we do not intend to dismiss a function of chirps in communication, we only intend to debate and discuss valid alternative hypothesis, advanced from reasonable considerations.

Before writing this manuscript, we have attempted to survey literally all the existing literature on chirps (including studies focused on behavior, peripheral sensory physiology as well as brain physiology). Although it is not unlikely that some studies have eluded our attention, an effort for a comprehensive review was made. Based on this survey we realized that none of the studies provided a clear and unambiguous piece of evidence to support the communication hypothesis (we refer here to the weak points highlighted in the discussion and mentioned in the previous comment). Which in fact does not come without its weak points and contradictions (see later comments).

It follows a summary of the mentions made to the communication theory in the different section of the manuscript including several edits we have applied in response to the Reviewer’s concern:

In the abstract we clearly state that we are considering an alternative that is only hypothetically complementary, not for sure. Nonetheless, we have identified a couple of instances that could sound dismissive of the “communication hypothesis” in the following section.

In the introduction we write in fact about the possibility of interference between communication signals and conspecific electrolocation cues, as they are both detected as beat perturbations. We did not mean to use “Interference” here as “reciprocal canceling”, rather we intended it as “partial or more or less conspicuous overlap” in the responses triggered in electroreceptors.

Hoping to convey a clearer message, we have edited the related statement and changed it to “both types of information are likely to overlap and interact in highly variable ways”.

We have also removed the statement: “According to this idea, beats and chirps are not only detected through the same input channel, but also used for the same purpose.” as at this point in the manuscript it may be too strong.

In the results section we do not include statements that might be seen as dismissive of the communication hypothesis but only statements in support of the “probing with chirps” idea (which is the central hypothesis of the study).

In the discussion paragraphs we elaborate on why the current functional view is either flawed or incomplete (first paragraph “existing functional hypotheses''). Namely: 1) multiple triggering factors implied in chirp responses covary and need to be disentangled (example DF/ sex), 2) findings on brown ghosts and a few other gymnotiforms have been used to advance the hypothesis of “communication through chirps'' in all weakly electric fish (including pulse species). 3) social encounters - in which chirps are recorded - imply also other behaviors (such as probing) which have not been considered so far. This point is related to the first one on covariates. 4) most studies referring to big chirps as courtship chirps were not done in reproductive animals (added now) and 5) no causal evidence has been provided so far to justify a role of chirps in social communication.

We are discussing these points as challenges to the communication hypothesis, not to dismiss the hypothesis, but rather to motivate future studies addressing these challenges.

We do not want to appear dismissive of the communication hypothesis and had therefore previously edited the manuscript to avoid the impression of exclusivity of the probing hypothesis. We have now gone over the manuscript once more and edited several sentences. Nevertheless, we want to point out again that - despite the large consensus - the communication hypothesis has, until now, never been investigated with the kind of rigor applied here.

The authors do acknowledge that chirps could function as both a communication and homeactive sensing signal, but it seems clear they wish to argue against the former and for the latter, and the evidence is not yet there to support this.

In both rounds of revision we have made an effort to convey a more inclusive interpretation of our findings. We tried our best to express our ideas as hypothetical, not as proof that communication through chirps does not exist. The aim of this study is to propose an alternative view, and this cannot be done without underlining the weak points of an existing hypothesis while providing and supporting reasonable arguments in favor of the alternative we advance. The actual evidence for a role of chirping in communication is much less strong than appears from the pure number of articles that have discussed chirps in this context.

Regarding the weak evidence against communication, here we can list a few additional important points related to the proposed interpretations of chirp function (more specific than those made earlier):

(1) A formally sound assessment of signal value/meaning - as typically done in animal communication studies should involve:

a) the isolation of a naturally occurring signal and determination of the context in which it is produced

b) the artificial replication of the signal

c) the observation that such mimic is capable of triggering reliable and stereotyped responses in a group of individuals (identified by sex and/or species) under the same conditions (conditioned, unconditioned, state-dependent, etc.). As discussed for instance in Bradbury and Vehrencamp, 2011; Laidre and Johnstone, 2013; Wyatt, 2015; Rutz et al., 2023.

This approach has so far not been applied to weakly electric fish. The initial purpose of the present study was in fact to conduct this type of validation.

(2) The hypothesis of chirps used for DF-sign discrimination - for “social purposes” - although plausible in the face of theoretical considerations, does not seem to be reasonable in practice, when one considers emission rates of 150 chirps per minute. We do find a strong correlation of chirp type with DF, which is often very abrupt and sudden (as if the fish were tracking beat frequency to guess its value) but the consideration made above on chirp rates seems to discourage this interpretation.

(3) The hypothesis of chirp-patterning (i.e. chirping may have meaning based on the sequence of chirps of different types, a bit like syllables in birdsongs) - assessed by only one study conducted in our group - has not been enough substantiated by replication. We have surveyed all possible combinations of chirps produced by interacting pairs in different behavioral conditions using different value for chirp sequence size: 2, 3,... ,8 chirps (both considering the sender alone as well as sender+receiver together). In all cases we found no evidence for a context dependent “modulation” of chirp types (i.e. no specific chirp type sequence in specific contexts).

(4) The hypothesized role of “large chirps” as courtship signals could be easily criticized by noting the symmetrical distribution of these events around a DF of 0 Hz . Although one could argue about a failure to discriminate DF-sign, to explain this well known pattern. However, we know from Walter Heiligenberg’s work and physiological considerations that such task can be solved easily through t-units and … in principle even just by motion (which would change the EOD phase in frequency dependent ways, thus potentially revealing the DF sign).

Overall, these considerations made us think that certainly chirping occurs in a social context, but it is the meaning of this behavior that remains elusive. We noticed that environmental factors are also strongly implied … we then formulate an alternative hypothesis to explain chirping but we do so without dismissing the communication idea.

All this seems to us just a careful way to critically discuss our results and those of other studies, without considering the issue resolved.

In the introduction, the authors state, "Since both chirps and positional parameters (such as size, orientation or motion) can only be detected as perturbations of the beat, and via the same electroreceptors, the inputs relaying both types of information are inevitably interfering." I disagree with this statement, which seems to be a key assumption. Both of these features certainly modulate the activity of electroreceptors, but that does not mean those modulations are ambiguous as to their source. You do not know whether the two types of modulations can be unambiguously decoded from electroreceptor afferent population activity.

We thank the Reviewer for noting this imprecision. We have addressed the Reviewer’s concern in another reply (see above).

My biggest issue with this manuscript is that it is much too strong in dismissing evidence that chirping correlates with context. In your behavioral observations, you found sex differences in chirping as well as differences between freely interacting and physically separated fish. Chirps tended to occur in close proximity to another fish. Your model of chirp variability found that environmental experience, social experience, and beat frequency (DF) are the most important factors explaining chirp variability. Are these not all considered behavioral or social context? Beat frequency (DF) in particular is heavily downplayed as being a part of "context" but it is a crucial part of the context, as it provides information about the identity of the fish you're interacting with. The authors show quite convincingly that the types of chirps produced do not vary with these contexts, but chirp rates do.

We believe the “perceived claim” may be an issue of unclear writing. We have now tried to better clarify that “context” affects chirp rates, but it does not affect chirp types as much (except when beat frequency is high).

We have edited two statements possibly susceptible to misinterpretation:

(1) In the results: “It also indicates that chirp parameters such as duration and FM do not seem to be associated with any particular context in a meaningful way, other than being affected by beat frequency.”

(2) In the discussion: the statement

“Recordings from interacting fish pairs confirmed the absence of any significant correlation between chirp type choice and behavioral context (Figure S2) although the variance of chirp parameters appears to be significantly affected by this factor (Figure 2). This may suggest that the effect of behavioral context is mainly detectable in the number of chirps produced (Figure S1), rather than the type (Figure S2).”

has been changed to:

“Recordings from interacting fish pairs confirmed the absence of any significant correlation between chirp type choice and behavioral context, except for those cases characterized by higher beat frequencies (Figure S2). This suggests that the effect of behavioral context highlighted in our factor analysis (Figure 2) is mainly due to the number of chirps produced (Figure S1), rather than their type (Figure S2).”

Eventually, in the results we emphasize the relatively higher impact of previously unexplored factors on chirp variance: “The plot of individual chirps (Figure 2C) shows the presence of clustering around different categorical variables and it reveals that experience levels or swimming conditions are important factors affecting chirp distribution (note for instance the large central “breeding” cluster in which fish are divided and the smaller ones in which fish are free). Sender or receiver identity does not individuate any clear clustering relative to either sex (see the overlap of male_s/male_r and female_s/female_r) or social status (dominant/subordinate). Chirps labeled based on tank experience (i.e. resident vs intruder) are instead clearly separated.”.

Further, in your playback experiments, fish responded differently to small vs. large DFs, males chirped more than females, type 2 chirps became more frequent throughout a playback, and rises tended to occur at the end of a playback. These are all examples of context-dependent behavior.

We do note that male brown ghosts chirp more than females. But we do also say - and show in figure 8 - that males move more in proximity to and around conspecifics. We do acknowledge that chirp time-course may be different during playbacks in a type-dependent manner. But how this can support the communication hypothesis - or other alternatives - is unclear. This result could equally imply the use of different chirp types for different probing needs. Since we cannot be sure about either, we do not want to put too much emphasis to it. Eventually, the fact that “context” (here meant broadly to define different experimental situations in which social but also physical and environmental parameters are altered) affects chirping is undeniable: cluttered and non-cluttered environments do represent different contexts which differently affect chirping in conspicuous ways.

In the results, the authors state, "Overall, the majority of chirps were produced by male subjects, in comparable amounts regardless of environmental experience (resident, intruder or equal; Figure S1A,C), social status (dominant or subordinate; Figure S1B) or social experience (novel or experienced; Figure S1D)." This is not what is shown in Figure S1. S1A shows clear differences between resident vs. intruder males, S1B shows clear differences between dominant vs. subordinate males, and S1D shows clear differences between naïve and experienced males. The analysis shown in Figure 2 would seem to support this. Indeed, the authors state, "Overall, this analysis indicated that environmental and social experience, together with beat frequency (DF) are the most important factors explaining chirp variability."

The Reviewer is right in pointing at this imprecise reference and we are grateful for spotting this incongruence. The writing refers probably to an earlier version of the figure in which data were grouped and analyzed differently. We now edited the text and changed it to: “Overall, the majority of chirps were produced by male subjects, at rates that seemed affected by environmental experience (resident, intruder or equal; Figure S1A,C), social status (dominant or subordinate; Figure S1B) and social experience (novel or experienced; Figure S1D).”

The choice of chirp type varied widely between individuals but was relatively consistent within individuals across trials of the same experiment. The authors interpret this to mean that chirping does not vary with internal state, but is it not likely that the internal states of individuals are stable under stable conditions, and that individuals may differ in these internal states across the same conditions? Stable differences in communication signals between individuals are frequently interpreted as reflecting differences between those individuals in certain characteristics, which are being communicated by these signals.

It seems here we have been unclear in the writing: while it is true that behavioral states are stable and can imply stable chirp patterning (if the two are related), since chirp types vary abruptly and in a reliable DF-dependent manner, different types of chirps are unlikely to be matched to different internal states following the same temporal order in such a reliable way (similarly repeated through consecutive trials).

This would imply the occurrence of different internal states in rapid sequence, reliably triggered by repeated EOD ramps, regardless of whether the playback is 20 sec long or 180 sec long.

We have edited this paragraph to better explain this: “The reliability by which the chirping response adapts to both the rate and direction of beat frequency is variable across individuals but rather stable across trials (relative to a given subject), further suggesting that chirp type variations may not reflect changes in internal states or in the animal motivation to specific behavioral displays (which are presumably subject to less abrupt variations and stereotypical patterning based on DF).”

I am not convinced of the conclusion drawn by the analysis of chirp transitions. The transition matrices show plenty of 1-2 and 2-1 transitions occurring.

The only groups in which 1-2 and 2-1 transitions are as frequent as 1-1 and 2-2 (being 1 and 2 the numerical IDs of the two interacting fish) are F-F pairs. This is a result of the fact that in females chirp rates are so low that within-fish-correlations end up being as low as between-fish-correlations. We believe the impression of the Reviewer could be due to the fact that these are normalized maps (see legend of Figure 5A-B).

Further, the cross-correlation analysis only shows that chirp timing between individuals is not phase-locked at these small timescales. It is entirely possible that chirp rates are correlated between interacting individuals, even if their precise timing is not.

We agree with the Reviewer, this is a possibility. To address this point, we did edit the results section to acknowledge that what we see may be related to the time window chosen (i.e. 4 sec):

“More importantly, they show that - at least in the social conditions analyzed here and within small-sized time windows - chirp time series produced by different fish during paired interactions are consistently independent of each other.”

Further, it is not clear to me how "transitions" were defined. The methods do not make this clear, and it is not clear to me how you can have zero chirp transitions between two individuals when those two individuals are both generating chirps throughout an interaction.

We thank the Reviewer for bringing up this unclear point. We have now clarified how transitions were calculated in the method section: “The number of chirp transitions present in each recording (dataset used for Figures 1, 2, 5) was measured by searching in a string array containing the 4 chirp types per fish pair, all their possible pairwise permutations (i.e. all possible permutations of 4+4=8 elements are: 1-1, 1-2, 1-3 … 7-6, 7-7, 7-8; considering the following legend 1 = fish1 type 1, 2 = fish 1 type 2, 3 = fish1 type 3 … 6 = fish2 type 2, 7 = fish2 type 3 and 8 = fish2 rise).”.

Zero transitions are possible if two fish (or groups of fish) do not produce chirps of all types. Only transitions of produced types can be counted.

In the results, "Although all chirp types were used during aggressive interactions, these seemed to be rather less frequent in the immediate surround of the chirps (Figure 6A)." A lack of precise temporal correlation on short timescales does not mean there is no association between the two behaviors. An increased rate of chirping during aggression is still a correlation between the two behaviors, even if chirps and specific aggressive behaviors are not tightly time-locked.

The Reviewer is right in pointing out the limited temporal scaling of our observations/analysis. We have now edited the last paragraph of the results related to figure 6 to include the possibility mentioned by the Reviewer: “The significantly higher extent of chirping during swimming and locomotion, consistently confirmed by 4 different approaches (PSTH, TM, CN, MDS), suggests that - although chirp-behavior correlations may exist at time-scales larger than those here considered - chirping may be linked more strongly with scanning and environmental exploration than with a particular motivational state, thus confirming findings from our playback experiments.”

The Reviewer here remarks an important point, yet, due to space limitations, we have considered only a sub-second scale. Most playback experiments in weakly electric fish implied the use of EOD mimics for a few tens of seconds - to avoid habituation in the fish behavioral responses - while inter-chirp intervals usually range between a few hundreds of milliseconds to seconds (depending on how often a fish would chirp). This suggested to us that a 4 second time window may not be a bad choice to start with.

In summary, it is simply too strong to say that chirping does not correlate with context, or to claim that there is convincing evidence arguing against a communication function of chirps. Importantly, however, this does not detract from your exciting and well-supported hypothesis that chirping functions in homeoactive sensing. A given EOD behavior could serve both communication and homeoactive sensing. I actually suspect this is quite common in electric fish (both gymnotiforms and mormyrids), and perhaps in other actively sensing species such as echolocating animals. The two are not mutually exclusive.

We agree with the Reviewer that context - broadly speaking - does affect chirping (as we mentioned above). We hope we have improved the writing and clarified that we do not dismiss communication functions of chirping, but we do lean towards electrolocation based on the considerations above made and our results.

We do conclude the manuscript remarking that communication and electrolocation are not mutually exclusive: ”probing cues could function simultaneously as proximity signals to signal presence, deter approaches, or coordinate behaviors like spawning, if properly timed (Henninger et al., 2018).” (see the conclusion paragraph of the discussion) .

Therein, we further add “These findings aim to stir the pot and initiate a discussion on possible alternative functions of chirps beyond their presumed communication role.”.

With this, we hope we’ve made it clear how we intend our manuscript to be read.

Reviewer #3 (Public Review):

Summary:

This important paper provides the best-to-date characterization of chirping in weakly electric fish using a large number of variables. These include environment (free vs divided fish, with or without clutter), breeding state, gender, intruder vs resident, social status, locomotion state and social and environmental experience, without and with playback experiments. It applies state-of-the-art methods for reducing the dimensionality of the data and finding patterns of correlation between different kinds of variables (factor analysis, K-means). The strength of the evidence, collated from a large number of trials with many controls, leads to the conclusion that the traditionally assumed communication function of chirps may be secondary to its role in environmental assessment and exploration that takes social context into account. Based on their extensive analyses, the authors suggest that chirps are mainly used as probes that help detect beats caused by other fish and as well as objects.

Strengths:

The work is based on completely novel recordings using interaction chambers. The amount of new data and associated analyses is simply staggering, and yet, well organized in presentation. The study further evaluates the electric field strength around a fish (via modelling with the boundary element method) and how its decay parallels the chirp rate, thereby relating the above variables to electric field geometry.

The main conclusions are that the lack of any significant behavioural correlates for chirping, and the lack of temporal patterning in chirp time series, cast doubt on a primary communication goal for most chirps. Rather, the key determinants of chirping are the difference frequency between two interacting conspecifics as well as individual subjects' environmental and social experience. The paper concludes that there is a lack of evidence for stereotyped temporal patterning of chirp time series, as well as of sender-receiver chirp transitions beyond the known increase in chirp frequency during an interaction.

These conclusions by themselves will be very useful to the field. They will also allow scientists working on other "communication" systems to perhaps reconsider and expand the goals of the probes used in those senses. A lot of data are summarized in this paper, with thorough referencing to past work.

The alternative hypotheses that arise from the work are that chirps are mainly used as environmental probes for better beat detection and processing and object localization, and in this sense are self-directed signals. This led to their prediction that environmental complexity ("clutter") should increase chirp rate, which is fact was revealed by their new experiments. The authors also argue that waveform EODs have less power across high spatial frequencies compared to pulse-type fish, with a resulting relatively impoverished power of resolution. Chirping in wave-type fish could temporarily compensate for the lower frequency resolution while still being able to resolve EOD perturbations with a good temporal definition (which pulse-type fish lack due to low pulse rates).

The authors also advance the interesting idea that the sinusoidal frequency modulations caused by chirps are the electric fish's solution to the minute (and undetectable by neural wetware) echo-delays available to it, due to the propagation of electric fields at the speed of light in water. The paper provides a number of experimental avenues to pursue in order to validate the non-communication role of chirps.

We thank the reviewer for the kind assessment.

Weaknesses:

My main criticism is that the alternative putative role for chirps as probe signals that optimize beat detection could be better developed. The paper could be clearer as to what that means precisely, especially since beating - and therefore detection of some aspects of beating due to the proximity of a conspecific - most often precedes chirping. One meaning the authors suggest, tentatively, is that the chirps could enhance electrosensory responses to the beat, for example by causing beat phase shifts that remediate blind spots in the electric field of view.

We agree with the Reviewer that a better and more detailed explanation of how beat processing for conspecific electrolocation may be positively affected by chirps would be important to provide. We are currently working on a follow-up manuscript in which we intend to include these aspects. For space limitations and readability we had to discard from the current manuscript a lot of results that could further clarify these issues.

A second criticism is that the study links the beat detection to underwater object localization. The paper does not significantly develop that line of thought given their data - the authors tread carefully here given the speculative aspect of this link. It is certainly possible that the image on the fish's body of an object in the environment will be slightly modified by introducing a chirp on the waveform, as this may enhance certain heterogeneities of the object in relation to its environment. The thrust of this argument derives mainly from the notion of Fourier analysis with pulse type fish EOD waveforms (see above, and radar theory more generally), where higher temporal frequencies in the beat waveform induced by the chirp will enable a better spatial resolution of objects. It remains to be seen whether experiments can show this to be significant.

Perhaps the Reviewer refers to the last discussion paragraph before the conclusions in which we mention the performance of pulse or wave-type EODs in electrolocation (referring here to ideas illustrated in a recent review by Crampton, 2019). We added to this paragraph a statement which could better clarify that we do not propose that chirping could enhance object electrolocation. What we mean is that, in a context in which object electrolocation occurs through wave-type EODs - given the generally lower performance of such narrow-band signals in resolving the spatial features of any object, even a 3D electric field - chirping could improve beat detection during social encounters by increasing the amount of information obtained by the fish.

The edited paragraph now reads: “While broadband pulse signals may be useful to capture highly complex environments rich in foliage, roots and other structures common in vegetation featuring the more superficial habitats in which pulse-type fish live, wave-type EODs may be a better choice in the relatively simpler river-bed environments in which many wave-type fish live (e.g., the benthic zone of deep river channels; Crampton, 2019). In this case, achieving a good spatial resolution is critical during social encounters, especially considering the limited utility of visual cues in these low-light conditions. In such habitats, social encounters may “electrically” be less “abrupt”, but spatially less “conspicuous” or blurred (as a 3D electric field may be). In such a scenario, chirps could serve as a means to supplement the spatial information acquired via the beat, accentuating these cues during periods of reduced resolution.”

Recommendations for the authors:

Reviewer #3 (Recommendations For The Authors):

None, my points in the original review have been properly addressed in this resubmission.

Changes from the earlier version to the revised version

Abstract

Abstract

Since the pioneering work by Moeller, Szabo, and Bullock, weakly electric fish have served as a valuable model for investigating spatial and social cognitive abilities in a vertebrate taxon usually less accessible than mammals or other terrestrial vertebrates. These fish, through their electric organ, generate low-intensity electric fields to navigate and interact with conspecifics, even in complete darkness. The brown ghost knifefish is one of the most widely studied species due to its rich electric vocabulary, made by individually variable and sex-specific electric signals. These are mainly characterized by brief frequency modulations of the oscillating dipole moment emitted continuously by their electric organ and are known as chirps. Different types of chirps are believed to convey specific and behaviorally salient information, serving as behavioral readouts for different internal states during behavioral observations. Despite the success of this model in neuroethology over the past seven decades, the code to decipher their electric communication remains unknown.

This study re-evaluates this view, aiming to offer an alternative, and possibly complementary, explanation for why these freshwater bottom dwellers emit electric chirps. By uncovering correlations among chirping, electric field geometry, and detectability in enriched environments, we present evidence for a previously unexplored role of chirps as specialized self-directed signals, enhancing conspecific electrolocation during social encounters.

Introduction

Homeoactive sensing (i.e. self-directed active signal production) differs from alloactive sensing in that energy changes are actively produced in order to affect object detection in the environment, as opposed to a simple reconfiguration of sensor parameters such as orientation or distance from the signal source (Zweifel and Hartmann, 2020). Homeoactive sensing is often used as a strategy to enhance spatial navigation in heterogeneous and unpredictable environments: by emitting high-pitch calls (8-200 kHz; Fenton et al., 1998), bats are able to detect features of stationary or moving objects while flying at high speeds and in complete darkness (Jones and Holderied, 2007). Similarly, odontocete cetaceans are capable of echolocating objects by emitting very short (a few ms long) clicks covering a broad frequency spectrum (5 Hz-100 kHz). A few other mammals (shrews and tenrecs), but also bird species (oilbirds and swiftlets) and even blind humans (Griffin, 1944; Thaler et al., 2011), are known to be capable of echolocating target objects through the emission of brief high-frequency echolocation calls (Konishi and Knudsen, 1979; Brinkløv et al., 2013; Gould, 1965). In general, all biosonar systems are good examples of homeoactive sensing strategies and represent a remarkable adaptation to low-light or turbid water environments in which other sensory modalities would offer only a fraction of the information required to optimize fast spatial navigation during foraging or social interactions.

But navigating in the dark is not only about detecting sounds. The electrosensory system of weakly electric fish is known to offer analogous advantages yet at the same time presenting fundamental differences. In the following study we focus on the brown ghost - one of the most studied species - and extend this analogy to the use of frequency modulations as useful signals in the context of conspecific electrolocation. Brown ghosts belong to the gymnotiform group (as does the electric eel) and occupy benthic habitats in the Amazon river basin (Crampton, 2019). They generate quasi-sinusoidal electric fields at relatively high frequencies (600-1000 Hz; Zakon et al., 2002) by means of modified motor neurons projecting mainly to the caudal region of the fish body where they form a so-called electric organ. For this reason, these electric fields are also referred to as electric organ discharges (EODs). EODs - and perturbations thereof - are detected through specialized electroreceptors distributed throughout the fish body and mainly concentrated around the head region. Because EODs are actively produced signals, this process is also called active electrolocation (as opposed to passive electrolocation, i.e. the sensing of weak bioelectric signals such as those emanating from muscle activity in prey or other organisms; Rose, 2004). Compared to echolocation, passive electrolocation is possible only at shorter ranges (0-5 cm; Lissmann and Machin, 1958; Von der Emde, 1999; Knudsen, 1974) and, although it is so fast to be considered almost instantaneous, it is more sensitive to signal degradation (Benda, 2020; Nelson and MacIver, 2006).

Active electrolocation allows to extend the active space of electrolocation (Hopkins and Westby, 1986; Brenowitz, 1982) slightly beyond such short range: since EODs spread roughly equally in all directions, all affected surrounding objects (preys, substrate, vegetation) are simultaneously detected (up to 10-12 cm away from the skin; von der Emde, 1999; Snyder et al., 2007). During social interactions, the range of active electrolocation can be further extended: when two fish approach each other, the alternating constructive and destructive interaction of their sinewave EODs results in an amplitude- and phase-modulated signal (also referred to as beat) detectable by each of them within a 30-40 cm radius (see appendix 1: Detecting beats at a distance; Knudsen, 1975; Fotowat et al., 2013; Henninger et al., 2018; Benda, 2020). The frequency of such beat modulations is equal to the frequency difference (DF) of the two sinewave signals. Because individual brown ghosts have each a different EOD frequency, detecting the beat can be used not only to sense an approaching conspecific but also to determine its identity and sex: at 27°C female EOD frequencies range between 600 and 800 Hz while male frequencies are usually between 800 and 1000 Hz. It follows that low DFs are indicative of same sex encounters, whereas high DFs are more typical of opposite sex interactions (Bastian et al., 2001).

In summary, weakly electric fish are capable of passively detecting bioelectric sources in close proximity to their body (0-5 cm), but also to actively sense inanimate objects as instantaneous EOD perturbations (up to 12 cm; Chen et al., 2005; Snyder et al., 2007; Fotowat et al., 2013). In addition, by detecting beats (i.e. the EOD amplitude modulations induced by other EOD emitting sources), they can instantaneously sense the presence of other conspecifics at distances of up to 30-50 cm (Fotowat et al., 2013; Henninger et al., 2018; Silva A., personal communication).

Importantly, the strength of electric fields decays with the cubed distance (see appendix 1) and EODs are subject to anisotropic degradation even at very short ranges. This is mainly due to the presence of noisy interferences encountered during locomotion through electrically heterogeneous underwater materials (Gómez-Sena et al., 2014; Pedraja et al., 2016; Milam et al., 2019; Yu et al., 2019).

The picture is further complicated by the presence of another critical source of beat interference: communication signals. These consist of brief EOD frequency modulations - or chirps - which can be broadly classified in at least 4 different types, based on their duration and the extent of the chirp frequency modulation (FM, Engler et al., 2000; Zupanc et al., 2006). Different types of chirps are thought to carry different semantic content based on their occurrence during either affiliative or agonistic encounters (this notion will be referred to as communication-hypothesis; Larimer and MacDonald 1968; Bullock, 1969; Hopkins 1974; Hagedorn and Heiligenberg 1985; Zupanc and Maler 1993; Engler et al. 2000; Engler and Zupanc 2001; Bastian et al., 2001).

Earlier version: Since both chirps and positional parameters (such as size, orientation or motion) can only be detected as perturbations of the beat (Petzold et al., 2016; Yu et al., 2012; Fotowat et al., 2013), and via the same electroreceptors, the inputs relaying both types of information are inevitably interfering. Moreover, as the majority of chirps are produced within a short range (< 50 cm; Zupanc et al., 2006; Hupé and Lewis 2008; Henninger et al., 2018; see appendix 1) this interference is likely to occur consistently during social interactions.

Revised version: Since both chirps and positional parameters (such as size, orientation or motion) can only be detected as perturbations of the beat (Petzold et al., 2016; Yu et al., 2012; Fotowat et al., 2013), and via the same electroreceptors, the inputs relaying both types of information are likely to overlap and interact in highly variable ways. Moreover, as the majority of chirps are produced within a short range (< 50 cm; Zupanc et al., 2006; Hupé and Lewis 2008; Henninger et al., 2018; see appendix 1) this interference is likely to occur consistently during social interactions.

Under the communication-hypothesis, the assumption that chirps and beats are conveying different types of information (i.e. semantic value as opposed to position and related geometrical parameters) is therefore leaving this issue unresolved.

Earlier version: In this study, we propose a solution to this problem by providing evidence for a previously unexplored function of chirping in weakly electric fish. We have gathered a large dataset of more than 70.000 chirps, obtained from staged social pairings, playback experiments and behavioral assessments of fish locomotory activity. We first show that the relative number of different chirp types in a given recording does not significantly correlate with any particular behavioral or social context. It is instead correlated to the biophysical properties of underwater electric fields and to the spatial arrangement of interacting EODs. By using high-speed infrared video recordings synchronized with multi-channel voltage data acquisition, we then analyzed chirps considering both their effects on beat processing and their spatial attributes, i.e. the spatial location of the interacting fish. Analyses of cross-correlation of chirp time-series, chirp-type transitions and of behavioral responses to playback chirps were used to assess the presence of meaningful correspondences. The occurrence of chirps of different types was then analyzed in relation with the most common behavioral displays used by brown ghosts during social interactions. Finally, by recording fish in different conditions of electrical “visibility”, we provide evidence supporting a previously neglected role of chirps: homeoactive sensing. According to this idea, beats and chirps are not only detected through the same input channel, but also used for the same purpose.

Revised version: In this study, we propose a solution to this problem by providing evidence for a previously unexplored function of chirping in weakly electric fish. We have gathered a large dataset of more than 70.000 chirps, obtained from staged social pairings, playback experiments and behavioral assessments of fish locomotory activity. We first show that the relative number of different chirp types in a given recording does not significantly correlate with any particular behavioral or social context. It is instead correlated to the biophysical properties of underwater electric fields and to the spatial arrangement of interacting EODs. By using high-speed infrared video recordings synchronized with multi-channel voltage data acquisition, we then analyzed chirps considering both their effects on beat processing and their spatial attributes, i.e. the spatial location of the interacting fish. Analyses of cross-correlation of chirp time-series, chirp-type transitions and of behavioral responses to playback chirps were used to assess the presence of meaningful correspondences. The occurrence of chirps of different types was then analyzed in relation with the most common behavioral displays used by brown ghosts during social interactions. Finally, by recording fish in different conditions of electrical “visibility”, we provide evidence supporting a previously neglected role of chirps: homeoactive sensing.

Results

Experience is the main factor affecting chirp variability

During social interactions, communication signals are often related to immediate aspects of behavioral physiology (e.g., internal states, Owren et al., 2010). As a consequence, correlations between signals and behavioral contexts are often indicative of signal meaning (Flack and De Waal, 2007). Thus, if behavioral meaning can be attributed to different types of chirps, as posed by the communication-hypothesis (e.g., Hagedorn and Heiligenberg, 1985; Larimer and MacDonald, 1968; Rose, 2004), one should be able to identify clear correlations between behavioral contexts characterizing different internal states and the relative amounts of different types of chirp.

To assess the presence of correlations between different behavioral contexts and the emission of different types of chirps, we recorded fish pairs in different settings based on 1) the experience with the recording environment (residents, intruders and equal, to define cases in which both fish of the pair were naïve to the environment), 2) the behavioral context including: dominance status (dominant and subordinate), simulated breeding conditions (breeding or no breeding), social experience (experienced, novel), and the direct physical accessibility of a conspecific (divided or free) for a total of N = 130 fish pairs (see methods for details).

EOD recordings took place in chambered aquariums using a multi-channel acquisition setup (Figure 1A). Chirp detection was carried out using a semi-automatic working pipeline (Figure 1B, see methods). Chirps have been previously categorized in different types based on the extent and the duration of their frequency modulation (FM in Hz and duration in ms; Engler et al., 2000). Initially, we searched for clusters in our chirp dataset using the k-means method and chirp FM and duration as input parameters (Figure 1C). Through this approach, 2 clusters were identified (clustering was validated using the silhouette method; Rouseeuw, 1987), suggesting a roughly bimodal distribution of chirps, according to these basic features (a third cluster could be defined by the separation between type 1 and type 2 chirps, which in our sample was not so obvious). Nonetheless, for consistency with the existing literature, we categorized chirps using rough cut-off values based on the distribution density of an early subset of chirps (N = 11342 chirps, see methods for details). These values roughly matched the broad parameter ranges previously reported for each type of chirp (Engler et al., 2000; Zupanc et al., 2006; Zakon et al., 2002; Bastian et al., 2001). Based on these cut-off values (duration = 50 ms, FM = 105 Hz), chirps were categorized in short (with small - type 2 - or large - type 1 - FM) and long (with small - rises - or large - type 3 - FM; Figure 1C). Slower events, such as gradual frequency changes (or gradual frequency rises; Engler et al., 2000), were more sporadic and were not considered in this analysis.

Chirp type categorization based on FM parameters.

A: Schematic diagram of the recording setup: holding aquariums (N = 8) were divided by a plastic mesh divider to prevent physical contact between the interacting fish (N = 2 per aquarium) while allowing electric interactions. Brown ghost EODs were detected by electrodes placed in each tank compartment (2 input channels per tank), amplified, digitized (20 kHz sampling rate) and recorded using custom-written MATLAB scripts. B: Electric signals and chirps were analyzed using their fast-fourier transforms (FFT, 1) and assigned to different fish based on the signal intensity (2). Frequency modulations were detected using a MATLAB-based heuristic method for automatic peak detection (APD, 3). False-positives, false-negatives or wrongly assigned chirps were revised manually (4) and chirp FM and duration were then measured (5). C: Chirp distribution by frequency modulation (Hz) and duration (ms). K-means chirp clustering indicated an optimal value of 2-3 for clustering chirps based on these two parameters (cluster centroids are indicated with black crosses; silhouette values are shown in the bottom right panel). The red lines indicate the cut-off values used to classify the different chirp types (type 1 = 1, type 2 = 2, type 3 = 3, rises = 4): duration 50 ms and modulation 105 Hz. These values are based on a dataset acquired at the beginning of the study (red, in the scatter plot; N chirps = 11342; N chirping fish = 16, N fish pairs = 8) and on the distribution density (D) of the whole chirp population (gray, in the scatter plot; N chirps = 30486; N chirping fish = 130, N fish pairs = 78). E: Representative examples of the 4 different chirp categories (voltage data on the first row, instantaneous frequency on the second and spectrogram on the third). F: Scatter plot showing the distribution of different chirp types by DF (frequency difference between sender and receiver fish): note the gradual change in chirp type composition (color coded) at different DF values (especially visible for type 1-3). Due to the sex difference in the brown ghost EOD frequency, negative DF values correspond mainly to females, positive values to males.

Earlier version: Overall, the majority of chirps were produced by male subjects, in comparable amounts regardless of environmental experience (resident, intruder or equal; Figure S1A,C), social status (dominant or subordinate; Figure S1B) or social experience (novel or experienced; Figure S1D). Interestingly, this sex difference was reversed when fish were allowed to interact freely, as opposed to being separated by a plastic mesh divider (Figure S1E,F). Sex differences were found also in the types of chirp produced, mainly in the case of male-to-female exchanges (MF; Figure S2C,D), confirming previous accounts of marked sex dimorphisms in both the number and the type of chirps produced by brown ghosts (Dye, 1987; Zupanc and Maler,1993; Dulka and Maler, 1994).

Revised version: Overall, the majority of chirps were produced by male subjects, at rates that seemed affected by environmental experience (resident, intruder or equal; Figure S1A,C), social status (dominant or subordinate; Figure S1B) and social experience (novel or experienced; Figure S1D). Interestingly, this sex difference was reversed when fish were allowed to interact freely, as opposed to being separated by a plastic mesh divider (Figure S1E,F). Sex differences were found also in the types of chirp produced, mainly in the case of male-to-female exchanges (MF; Figure S2C,D), confirming previous accounts of marked sex dimorphisms in both the number and the type of chirps produced by brown ghosts (Dye, 1987; Zupanc and Maler,1993; Dulka and Maler, 1994).

To better define the main factors underlying chirp variability (i.e. types, frequency of beat and carrier signals, behavioral context, etc.), we used a 3D factor analysis for mixed categorical and numerical data (FAMD) and evaluated for each recorded chirp several parameters relative to the interacting fish (both sender and receiver) and their respective EODs (see methods for details; Figure 2). The model obtained was adjusted in order to include the most significant variables (those explaining above 20% of the total variance). The lowest number of variables required to explain the highest percentage of variance (N = 3: tank experience, DF and context) was obtained by further simplifying the model through subsequent iterations (percent of variance explained 52.3%).

Factor analysis of mixed data (FAMD) – social contexts and tank experience.

A: Scatter plot showing the contribution of all chirp-related variables to the overall variance of the whole chirp dataset: among these, EOD parameters such as amplitude (EODamp), frequency (EODs or EODr, based on sender or receiver identity), spectral power density (pow) were considered together with variables related to chirps, such as frequency modulation (freq. modulation), duration, sex of sender or receiver fish (sex_s, sex_r), the time of occurrence within a 1 hour trial (timestamp), the type (1 to 4), the DF. Variables related to the fish experience with either the tank environment (tank experience: resident = 1-week tank experience, intruder = new to the tank, equal = both new) and experience with the paired conspecific (context). This latter category refers to the reciprocal experience of each fish pair (novel or experienced), their hierarchical status (dominant or subordinate), the type of interaction (divided = behind a plastic mesh barrier, free = freely swimming) and the simulated breeding season (based on water conductivity levels: high conductivity = ca. 400 μS, no breeding; low conductivity = ca. 100 μS, breeding) at which the interaction takes place (see methods for details). Triangles indicate the coordinates of the variable centroids, their contribution (“contrib”) is coded by color intensity, whereas the quality of their representation on the transformed coordinates is coded by color hue (“cos2”). B: Estimates of the total variance explained indicate that tank experience, together with DF and context, are the most important factors explaining chirp variance. C: Representation of chirps in the transformed coordinates. The clustering is based on qualitative coordinates (tank experience, context and chirp type). Cluster distance represents the correlation among variables. The marginal plots show the kernel distribution of the chirp population color-coded according to chirp type (legend on the bottom right). Labeling chirps by DF shows how chirps can be meaningfully clustered based on this parameter (inset, top right).

Earlier version: Overall, this analysis indicated that environmental and social experience, together with beat frequency (DF) are the most important factors explaining chirp variability (Figure 2B). It also indicates that chirp parameters such as duration and FM do not seem to be associated with any particular context in a meaningful way. The plot of individual chirps (Figure 2C) shows the presence of clustering around different categorical variables and it reveals that experience levels or swimming conditions are important factors affecting chirp distribution (note for instance the large central “breeding” cluster in which fish are divided and the smaller ones in which fish are free). Sender or receiver identity does not individuate any clear clustering relative to either sex (see the overlap of male_s/male_r and female_s/female_r) or social status (dominant/subordinate). Chirps labeled based on tank experience (i.e. resident vs intruder) are instead clearly separated. Social experience may also be relevant although it appears to affect chirp clustering only slightly (note the contiguity of the novel and experienced clusters). Nonetheless, most of the clustering could be simply explained by the DF (see inset in Figure 2C), as suggested by the evaluation of chirp types in the social contexts considered (Figure S2).

Revised version: Overall, this analysis indicated that environmental and social experience, together with beat frequency (DF) are the most important factors explaining chirp variability (Figure 2B). It also indicates that chirp parameters such as duration and FM do not seem to be associated with any particular context in a meaningful way, other than being affected by beat frequency. The plot of individual chirps (Figure 2C) shows the presence of clustering around different categorical variables and it reveals that experience levels or swimming conditions are important factors affecting chirp distribution (note for instance the large central “breeding” cluster in which fish are divided and the smaller ones in which fish are free). Sender or receiver identity does not individuate any clear clustering relative to either sex (see the overlap of male_s/male_r and female_s/female_r) or social status (dominant/subordinate). Chirps labeled based on tank experience (i.e. resident vs intruder) are instead clearly separated. Social experience may also be relevant although it appears to affect chirp clustering only slightly (note the contiguity of the novel and experienced clusters). Nonetheless, most of the clustering could be simply explained by the DF (see inset in Figure 2C), as suggested by the lack of association of chirp types to the other social variables considered (Figure S2).

Brown ghosts display invariant chirp responses to chirping playback EOD mimics

A complementary approach to uncover correspondences between communication signals and behavioral states relies on the use of playback experiments. To determine whether brown ghosts respond to playback chirps in a chirp-type specific manner, we adopted a variation of the classic “chirp chamber” approach (e.g. Dye, 1987; Bastian, 2001; Dunlap and Larkins-Ford 2003) and recorded brown ghosts’ electrical and locomotor responses to playback EODs in freely moving conditions (N = 16 fish, 8 females, 8 males; Figure 3A). In these experiments, the matching of emitted signals to playbacks as well as the locomotor activity (video recorded) displayed during close-range interactions with EOD sources (i.e. decisions to approach or avoid playback signals), can be used as a possible measure of signal value (Fugère and Krahe, 2010; King and McGregor, 2016).

Invariant chirping responses to playback chirps in freely swimming fish.

A: Schematic diagram of the setup used for the playback experiments: both the fish EOD and the fish swimming behavior are recorded during 60 s long playback trials. During each trial the fish locomotion is scored based on the % coverage of tank space (60 x 30 cm) in 4 regions of interest (ROI) at increasing distance from the playback electrodes (1 = close, 2-3 = intermediate, 4 = far). Playback trials are organized in 4 different modes (0-3), each including 15 DF levels, all shuffled and randomized. B: The box plots on the top row show the total number of chirps produced by either male or female brown ghosts produced in response to plain EOD mimics (mode 0), sinewaves containing type 3 chirps (red, mode 1), sine waves containing trains of type 2 chirps (blue, mode 2) or rises (grey, mode 3). Chirp counts relative to each individual fish are summed across different DFs. The boxplots on the bottom row show the trial scores relative to the same subjects and summed across different DFs. C: Heatmaps showing the DF-dependent but mode-invariant distribution of mean chirp types produced by male and female subjects in response to different playback regimes. D: Score heatmaps showing that fish of both sexes approach playback sources (i.e. higher scores in ROI 1) with equal probability regardless of playback types or DFs. Female fish may be more stationary in proximity of the electrodes, resulting in slightly higher ROI-1 scores.

Fish were allowed to swim in a 30 x 80 cm aquarium at whose opposite ends two 13 x 30 cm playback compartments were created using plastic mesh barriers. Playback electrodes were positioned in either compartment (randomly chosen) and parallel to the longest tank axis, ca. 10 cm apart. Electric stimuli were EOD mimics at different frequency steps (DFs) of both positive and negative sign, relative to the fish’s own EOD frequency (values in Hz: ±240, ±160, ±80, ±40, ±20, ±10, ±5, 0). The EOD playback intensity was set at 1mV/cm at ca. 5 cm distance (which is in the range of natural field strengths of this species). Playback sessions were organized in 4 different modes, based on the pattern of frequency modulations contained (mode 0 = no chirps, mode 1 = type 3 chirps, mode 2 = type 2 chirps, mode 3 = rises). Chirp parameters were set to represent rises and the 2 main chirp categories defined by the k-means clustering without the constraints of our cut-off values (short - type 1 and type 2 - and long - type 3 - chirps; Figure 1C,D). Each playback trial lasted ca. 60 sec with 180-sec inter-trial intervals (see methods for details). An infra-red video camera synched with the EOD recording was used to score the average fish position during trials (see methods).

Fish responses to plain sinewave EODs confirmed previously reported findings: fish respond with large chirps (type 1 and type 3) to high DFs and with small chirps (type 2) to low DF signals (Figure S3), with males producing more chirps than females (Figure 3B; Bastian et al., 2001; Engler and Zupanc, 2001). Notably, a transient decrease in chirping was observed in the immediate temporal surround of playback chirps (modes 1-3, Figure S4). This effect was similar for all chirp types and seemed to be limited to a short time around chirp occurrence. Accordingly, a more pronounced decrease in chirping was observed during 3 Hz trains of type 2 chirps, possibly due to the higher frequency at which these events were repeated (Figure S4). Apart from this, fish responses to chirping EODs were hardly distinguishable from non-chirping EOD mimics in terms of number of chirps produced (Figure 3C,D; but see also Triefenbach et al., 2008).

Analysis of the time course of chirping further revealed two interesting details: first, while type 1 and type 3 chirps seem to be randomly scattered throughout the whole trial (60 s), the number of type 2 chirps seems to increase over time (Figure S5). This trend was independent of the type of playback and is therefore intrinsically related to the use of this particular type of chirp. All chirp types are subject to a brief decrease only in concurrence with 2 Hz trains of type 2-chirp playback (Figure S5). Second, the vast majority of rises was produced right at the end of the playback signal, regardless of whether or not playback chirps were occurring (Figure S5).

Overall, these results show that brown ghosts do not respond to playback chirps in a chirp-type specific manner. They are coherent with previous reports showing an increase of type 2 chirps over time during paired interactions (Hupé and Lewis 2008) and the presence of a short delay between consecutive chirps produced by different subjects (Hupé and Lewis 2008; Henninger et al., 2018). The consistent occurrence of rises at the end of playback signals has also been reported previously (Kolodziejski et al., 2007). Eventually, the strong symmetrical dependency of chirp production on positive and negative beat frequencies (DFs; Figure 3C, Figure S3; Bastian et al., 2001; Engler and Zupanc, 2001; Triefenbach and Zakon, 2007; Kolodziejski et al., 2007) and the lack of specific matching in the response of brown ghosts to playback chirps confirms that, in this species, chirping is primarily conditioned by beat parameters (Walz et al., 2013). Finally, these findings indicate that, irrespective of the meaning or valence chirps may have, the production of chirps and the detection of non self-generated chirps may be two independent processes.

Eventually, if chirping depends merely on beat parameters (such as beat frequency, i.e. DF, or contrast) one would expect to observe invariant responses to stimuli of different duration as well as across repeated trials, because the emission pattern of chirps of different types would adjust to frequency changes in similar ways across trials.

To evaluate this possibility, we stimulated a group of brown ghosts with playback signals with frequency ramps (DF +/- 300 Hz, in both directions). Stimuli were provided with different durations (20, 60, 180 sec), each in 3-4 repeats. Individual fish were caged in a mesh tube and playback stimuli were calibrated as previously done.

Earlier version: Perhaps not surprisingly, we found that chirps are produced at higher rates when playback frequency gets closer to the fish’s own EOD frequency (i.e. reaching a low DF value), similarly to what happens when static frequency playbacks are used. Yet, the choice of chirp type seemed to be highly variable across individuals (Figure 4D-F). Despite this heterogeneity, each individual fish displayed rather similar chirping patterns, both across trials of different length as well as in repeats of the same trial (Figure 4D-F). In addition, regardless of the type of chirp produced in a given trial, the overall number of chirps was reliably correlated with trial duration with an astonishingly high linear correlation coefficient (R2 = 0.987, Figure S6). The relative number of the different types of chirps also varied depending on playback duration, possibly as an effect of ramping speed (pie charts in Figure 4G-I). As previously shown, chirp type choice on average appears to reflect changes in the DF sign: larger chirps are scarcer around 0 Hz, while smaller ones are more frequent around the smaller DF values (Figure 4G-I). In some cases, this implied a quick DF-dependent transition from large to small chirps (see fish IDs S6, G5 and P3 for instance, Figure 4D-F). Nonetheless, our results show a general trend for all chirp types to be produced more often around lower DF values, as evident by looking at inter-chirp intervals (ICIs, Figure S7). The reliability by which the chirping response adapts to both the rate and direction of beat frequency is variable across individuals, further suggesting that chirp type variations may not reflect changes in internal states or in the animal motivation to specific behavioral displays (presumably more conserved in the tested subjects). Instead, the choice of different chirp types may reflect the presence of different sensory demands posed by different beat frequencies to the processing of the beat itself (Barayeu et al., 2023). Taken together, these results indicate that chirp patterning is mainly affected by beat frequency (i.e. DF).

Revised version: Perhaps not surprisingly, we found that chirps are produced at higher rates when playback frequency gets closer to the fish’s own EOD frequency (i.e. reaching a low DF value), similarly to what happens when static frequency playbacks are used. Yet, the choice of chirp type seemed to be highly variable across individuals (Figure 4D-F). Despite this heterogeneity, each individual fish displayed rather similar chirping patterns, both across trials of different length as well as in repeats of the same trial (Figure 4D-F). In addition, regardless of the type of chirp produced in a given trial, the overall number of chirps was reliably correlated with trial duration with an astonishingly high linear correlation coefficient (R2 = 0.987, Figure S6). The relative number of the different types of chirps also varied depending on playback duration, possibly as an effect of ramping speed (pie charts in Figure 4G-I). As previously shown, chirp type choice on average appears to reflect changes in the DF sign: larger chirps are scarcer around 0 Hz, while smaller ones are more frequent around the smaller DF values (Figure 4G-I). In some cases, this implied a quick DF-dependent transition from large to small chirps (see fish IDs S6, G5 and P3 for instance, Figure 4D-F). Nonetheless, our results show a general trend for all chirp types to be produced more often around lower DF values, as evident by looking at inter-chirp intervals (ICIs, Figure S7). The reliability by which the chirping response adapts to both the rate and direction of beat frequency is variable across individuals but rather stable across trials (relative to a given subject), further suggesting that chirp type variations may not reflect changes in internal states or in the animal motivation to specific behavioral displays (presumably subject to less abrupt variations and stereotypical patterning based on DF). Instead, the choice of different chirp types may reflect the presence of different sensory demands posed by different beat frequencies to the processing of the beat itself (Barayeu et al., 2023). Taken together, these results indicate that chirp patterning is mainly affected by beat frequency (i.e. DF).

Responses to EOD frequency ramps.

A-C) Examples of spectrograms from playback trials showing the responses to frequency ramps (increasing from -300 Hz to 300 Hz or decreasing within the same range) of different durations: 20, 60 and 180 sec respectively. D-F) Chirp raster plots relative to the three trial types. Chirp responses are grouped based on fish identity (unspaced rows represent responses by the same fish). Playback time is represented by the gray areas. Some fish produced chirps even in absence of a playback EOD. G-I) Histograms showing chirp type distributions by DF. Trials in which decreasing ramps were used were adjusted by flipping the time array to match the DF values. The pie plots show the relative amounts of the four types of chirp.

Chirping is mainly a self-referenced behavior

As shown in other communication systems (bats, Bohn et al., 2008; cetaceans, Dunlop, 2017; primates, Slocombe and Zuberbüler, 2005a,b; Notman and Rendall, 2005), valence and behavioral meaning could reside in conserved motifs that can serve as unequivocal signifiers for the specific context in which they are produced (Berwick et al., 2011; Nieder and Mooney, 2020). Similarly, in brown ghosts, while single chirps may neither be behaviorally salient nor informative, meaning could be attributed to chirp patterns or chirp transitions produced by the senders of the signals, the receivers or both. In agreement with this idea, in weakly electric fish, specific chirp sequences have been proposed to be relevant to specific behavioral contexts (Henninger et al., 2018; Dye, 1987) and even emotional states (Triefenbach and Zakon, 2008; Smith et al., 2013; Silva et al., 2013).

To validate this idea and assess whether specific chirp patterns could be considered signatures for different types of behavioral interactions, we evaluated all possible chirp sequences in the chirp time-series recorded in our first set of experiments (Figure 1). This was done first independently (considering each subject separately; Figure S2B,D) and then interactively (considering sender-receiver transitions; Figure 5).

Chirp time-series correlations during social interactions.

A,B: Chirp transition maps representing the median transition probability (normalized) of all possible chirp type pairs calculated for female and male fish, considered together. For each type of social pairing, the identity of the two fish is indicated by different ID numbers in the different map quadrants: 1-1 = fish 1 to fish 1, 1-2 = fish 1 to fish 2, 2-1 = fish 2 to fish 1, 2-2 = fish 2 to fish 2. In mixed pairs, the order given in the sex tag of each plot follows the same order of the fish ID. The presence of higher chirp-transition frequencies in the 2nd and 4th quadrants of the matrices (labeled with 1-1 and 2-2) indicates a substantial independence between chirp time-series (i.e. lack of temporal correlation). Transitions from female to male fish (A) or male-to-female fish (B) are considered separately. For each case, chirps were selected based on the sex of the sender fish. The 1-2 numbering refers to fish identity. On the right side of each matrix, chirp totals are displayed in boxplots for sender and receiver fish (outliers are represented as dots lying beyond the boxplot whiskers; red bars = 100) and cross-correlation indices (cci, 50 ms binning) are provided for chirp time series relative to the same fish pairs (red dotted lines = confidence intervals corresponding to 3 cci standard deviations). C,D: Chirp transitions relative to same sex pairings. A higher level of interaction for F-F pairs (visible in the first and third quadrants in C) is probably due to the extremely low chirp rates in these pairs. Notably, higher chirp rates (as in M-M pairs, D,E) do not result in higher cross-correlation levels. E,F: Chirp transitions for resident-intruder pairs (M-M and F-F). Since most chirps are produced by M-M resident-intruder pairs in divided aquariums, plots in E resemble those in D, as they are relative to overlapping datasets. G,H: Chirp transitions for dominant-subordinate pairs (M-M and F-F). I,J: Chirp transitions for freely swimming (naive, I or experienced, J) opposite-sex pairs. Note the reversed sexual dimorphism in chirp rates, in both cases.

Since this analysis led to very similar results for different chirp sequence lengths (2, 4, 6, 8 chirps), for simplicity we focus here on the simplest case of the pairwise transitions.

Considering all recordings independently, type 2 chirps are overall the most frequently produced chirp type, both in females as well as males (Figure S2A,C). As a direct consequence, the most common chirp type transition found is “type 2”-”type 2”, except for the case in which males are chirping to females (MF; Figure S2B). In this case, repeats of type 1 chirps become more prominent (especially in experienced pairs), possibly due to the higher amount of type 1 chirps produced by males at high frequency beats (i.e. high DF; Figure 1F, Figure S2C) and to the concurrent decrease in the number of type 2 chirps produced by experienced fish (see the effect of chirp timestamp in Figure 2A and experience in Figure S1D). In any other case, neither chirp transitions nor chirp type ratios were differently distributed. Rather, the proportion of different chirp types and the transition frequencies from each type to the others were remarkably similar, regardless of the context and the sex of the interacting fish (Figure S2). Overall, these results show that, at an individual level, the most common chirp sequence is a repeat of type 2 chirps in both males and females, with type composition being only slightly more heterogeneous for male senders, when paired with female receivers. This further supports the idea that chirp type choice relies mainly on beat frequency (DF). In other words, at least when chirp sequences are considered individually, chirp-type patterning is highly conservative and not connected to any of the behavioral contexts considered, except the one identified by low vs high DFs.

Yet, it is possible that meaningful chirp transitions (i.e. correlated to a specific behavioral context) are to be identified only when considering time series interactively (Chow et al., 2015; Fröhlich, 2017). Recordings from brown ghosts, conducted both in the field and in captivity, suggested the presence of correlations between male chirp rate (type 2 and type 3) and female larger chirps (type 3; Henninger et al., 2018). Meaningful correlations have been also suggested to consist more broadly in temporal correlations between time-series, with chirps produced with a preferred latency by two interacting subjects. However, different authors reported different latency values (Zupanc et al., 2006, latency 500-1000 ms; Hupé and Lewis 2008, latency 200-600 ms; Henninger et al., 2018, latency 165 ms). These discrepancies could be due to either breed differences or differences in the recording settings.

In our recordings, although the temporal correlation of chirp series was highly variable among fish pairs (Figure 5), it consistently depended on the number of chirps produced (Zupanc et al., 2006, estimated maximum chirp rate 0.4/sec; Hupé and Lewis 2008, estimated maximum chirp rate 0.53/sec; Henninger et al., 2018, estimated maximum chirp rate 3/sec) raising the question of whether the previously reported instances of higher temporal correlations could be just a direct consequence of higher chirp rates occurring when fish get closer to each other. We assessed chirp transitions considering both senders and receivers interactively for different sex pairings (M-M, M-F, F-M, F-F; Figure 5A-D). Transition frequencies were mapped onto four quadrants defined by fish identity (1 or 2) and temporal order: 1-1, 1-2, 2-1 and 2-2.

What stood out most consistently from this analysis was the much lower (or sometimes completely absent) extent of sender-receiver chirp transitions, as shown by the higher transition probabilities along the 1-1 and 2-2 quadrants of the probability maps. According to this, cross-correlation indices are very low or near zero levels in all cases (Figure 5). Although slightly different interaction rates were observed in different contexts, in all cases the most frequent chirp transitions were self-referenced (1-1 or 2-2). Although within FF pairs 2-1 and 1-2 interactions seem to be more frequent (Figure 5C,F,H), this is probably due to the much lower number of chirps produced by this sex (see boxplots showing the total number of chirps in each group).

Eventually, the absence of conserved temporal correlations could also depend on the type of behavioral interactions and the proximity at which they take place (Dunlap, 2002). Social distance in particular could explain the high chirp-rates previously observed during courtship (Hagedorn and Heiligenberg 1985; Henninger et al., 2018) and short-range agonistic interactions between fish confined in a relatively small aquarium (Hupé and Lewis, 2008) as opposed to the more sporadic chirping observed during encounters over wider spatial ranges (Henninger et al., 2018). Alternatively, the higher chirp temporal cross-correlation previously observed by other authors could be due to neuro-endocrine effects of courtship or other mating-related factors.

Earlier version: To test whether cross-correlation would be significantly affected during courtship, we simulated the onset of the reproductive season by gradually lowering water conductivity over the course of a 40-day period (from 400 μS to 100 μS, Kirschbaum, 1979). As a result of this treatment, at least 75% of the females used in this experiment (N = 8) showed visible signs of egg maturation. Mixed-sex pairs (N = 8 pairs) held in such conditions were then allowed to interact freely, in absence of any tank divider. As a control group we paired males and females in absence of any water conductivity change (“novel” vs “experienced”, Figure 5I,J). Although female fish produced a surprisingly higher number of chirps in these conditions, only minimal signs of cross-correlation could be detected, without reaching on average the threshold for statistical significance (Figure 5J). Collectively, these results indicate that brown ghosts utilize a very limited variety of chirp type transitions in different behavioral contexts and do not emit specific chirp patterns in any given behavioral context analyzed. More importantly, they show that - at least in the social conditions analyzed here - chirp time series produced by different fish during paired interactions are consistently independent of each other. This further confirms the notion of chirp production and detection as independent processes and that chirp type variability merely depends on beat frequency (DF).

Revised version: To test whether cross-correlation would be significantly affected during courtship, we simulated the onset of the reproductive season by gradually lowering water conductivity over the course of a 40-day period (from 400 μS to 100 μS, Kirschbaum, 1979). As a result of this treatment, at least 75% of the females used in this experiment (N = 8) showed visible signs of egg maturation. Mixed-sex pairs (N = 8 pairs) held in such conditions were then allowed to interact freely, in absence of any tank divider. As a control group we paired males and females in absence of any water conductivity change (“novel” vs “experienced”, Figure 5I,J). Although female fish produced a surprisingly higher number of chirps in these conditions, only minimal signs of cross-correlation could be detected, without reaching on average the threshold for statistical significance (Figure 5J). Collectively, these results indicate that brown ghosts utilize a very limited variety of chirp type transitions in different behavioral contexts and do not emit specific chirp patterns in any given behavioral context analyzed. More importantly, they show that - at least in the social conditions analyzed here and within small-sized time windows - chirp time series produced by different fish during paired interactions are consistently independent of each other. This further confirms the notion of chirp production and detection as independent processes and that chirp type variability merely depends on beat frequency (DF).

Chirping correlates with active sensing behaviors

To gain better insights into the behavioral correlates of chirping, we conducted simultaneous video and EOD recordings of freely swimming fish pairs (N = 12, mixed sexes; see methods). The aim of these experiments was to assess the occurrence of consistent chirp-behavior correlations and more in general to evaluate chirp patterning in conjunction with behaviors other than chirp production. Behaving fish were recorded using an infrared camera under IR illumination (920 nm). Fish were allowed to habituate to the testing conditions (15 min) prior to recording of both behavioral and electrical activity (15 min). In this phase, fish were allowed to sense each other electrically through a plastic mesh divider. EODs were recorded in this phase and after removing the mesh barrier. Of the 23 fish chirping during these trials, 12 increased chirping once the direct interaction started, 5 did not alter their chirping rates, 6 decreased them. Chirps were coded again in 4 types. The ethogram used for this analysis was created according to previously published observations (Triefenbach and Zakon, 2008) and included the following behaviors: attacks (i.e. bites to the fish head, atk), chasing attempts (ch), escape attempts (esc), head butts (i.e. direct hits of the snout to the target fish flank, head), jaw quivering (often preceding attacks directed to the head or the mouth, jaw), backward swimming (knife), forward locomotion (loc), stationary motion or inactivity (rest), backward approaches with tail contact (tail1) or involving tail waving (tail2), and close body contact (not always aggressive in nature, accompanied by side-to-side swimming, body intertwining and even periods of inactivity, wrestle).

Peristimulus time histograms (PSTH) calculated around the annotated chirps (N = 4672 chirps, of which 252 were type 1, 4164 were type 2, 77 were type 3 and 179 were rises) showed no behavior reliably occurring immediately after or before chirping (time window = 4 sec). Chirps occurred often during active locomotion (directed either forward or backwards), type 3 chirps seemed to occur more evidently towards the end of a forward-oriented swimming bout, while rises appeared to be centered around forward locomotion (loc, Figure 6A). Although all chirp types were used during aggressive interactions, these seemed to be rather less frequent in the immediate surround of the chirps (Figure 6A).

Chirping is mainly correlated with locomotion and active sensing.

A: Peri-stimulus time histograms (PSTH) centered around different chirp types (window = 4 sec) indicating that chirps are reliably emitted during locomotion-related behaviors (“knife”, “loc”, and “tail1”) but also during resting states (“rest”). Chirps do not seem to trigger any defensive or aggressive behavior directly (i.e. no evident left/right bias in the PSTH). A more detailed description of each behavior is provided in the methods section (see also panel C). B: Transition matrix showing the total number of behavioral transitions (data pooled from all pairs). Note that all behaviors are considered both individually as well as together with those co-occurring (example: knife-type2 indicates a type 2 chirp produced during backward swimming). The color threshold has been lowered to visualize the most significant transitions. The marginal bar plots represent the sum of the number of transitions along each axis. C: To extract meaningful behavioral correlations, a text-mining algorithm (see methods) was used to perform a co-occurrence analysis on the whole set of annotations in chronological order (data pooled from 12 fish pairs). In the graph, words closely associated with each other are linked with darker lines (according to a 0-1 coefficient scale). This analysis emphasizes the presence of modularity in a group of strings (i.e. chirps and behaviors), which represents the degree of partitioning of the whole word dataset (N = 5). Each cluster delineates strings sharing similar co-occurrence patterns (see color code). D: The multi-dimensional scaling of the same word-database (2D in this case) shows behavioral combinations having similar patterns of occurrence, regardless of their co-occurrence with other behavioral instances. For this representation, each word is considered individually and the color code represents the clustering of objects based on their proximity (method: Kruskal, distance: Jaccard, N of clusters = 8).

A tendency for a decrease in attack rates (atk and head) could be seen perhaps more clearly for type 2 chirps, possibly due to the higher number of this chirp type (89% of all chirps were type 2). Chirps and behaviors occurred with some overlap, depending on whether they represented focal instances or continuous states (such as swimming, wrestling or resting). To assess the presence of direct correlations between individual and composite behavioral patterns including chirps, aggressive, defensive or explorative behaviors, a transition matrix (TM) was calculated using all recorded transitions (Figure 6B).

This allowed us to make the following considerations: 1) as previously observed, type 2 chirps are the most represented chirp type. They are mostly produced in repeats and predominantly during backward swimming (knife), they are also followed - to a lower extent - by direct contact (wrestle), head bites (atk), backward tail-to-tail swimming (tail1), escape attempts (esc), type 3 chirps, head butts and chasing; 2) behaviors occurring during aggressive encounters (offensive or defensive) are related to one another and also occur in repeats; 3) most chirps occurring during backward swimming are followed by other chirps within the same activity bout, but if attacks are present (knife + … + atk or +head) chirps are not (in other words attacks may follow chirps but not during backward motion); 4) when used during forward swimming, chirps are often preceding a chasing attempt or a reversal in the swimming direction (loc knife).

While this approach allows to assess chirp-behavior correlations considering even complex combinations, the occurrence of over-represented instances may bias the results by masking potentially significant but less frequent correspondences between complex behavioral patterns. In addition, it has the potential limitation of exaggerating differences among similar behavioral patterns simply because they do not always occur in an identical way.

A way to overcome these issues is to conduct a nearest-neighbor analysis of behavioral sequences, considering not only strictly consecutive instances (as typically done in a transition matrix) but also more general probabilistic correlations among them. To this aim, we have used a quantitative content analysis approach using freely available software for text mining (KH Coder; Higuchi and Koichi 2016; Figure 6C-D). Notably, this analysis handles words only from a statistical point of view (i.e. as character strings) and does not evaluate semantic correlations. Data from all behavioral observations were pooled into one single text file, maintaining the original chronological order, subsequently uploaded in KH-C. We first calculated a co-occurrence network (CON), which linked different strings based on their co-occurrence within the dataset. Importantly, this calculation does not consider only pairwise transitions in which a given word is used but also deals with whole sets of correlated transitions. As a result, words are clustered together based on similarities in the usage pattern of their transition sets (Higuchi and Koichi 2016).

The resulting co-occurrence diagram confirmed that chirps are used more often in association with swimming activity (knife, loc, tail1, esc) and direct physical interaction between fish (wrestle). More explicitly aggressive behaviors remained more loosely connected (Figure 6C). Interestingly, resting states and rises seem to be used in similar ways (Figure 6C) despite the scarce transition probability from one to the other. Conversely, forward swimming (loc) and chasing (ch) behaviors seemed to co-occur quite often.

To validate these results, we also conducted a factor analysis (multidimensional scaling, MDS) to represent different behaviors on transformed coordinates based on their ranking and usage in the dataset (Figure 6D). This complementary approach was used to evaluate more simply the extent of clustering among words having similar appearance patterns, regardless of co-occurrence frequencies. Factor analysis confirmed that type 2 chirps are mostly associated with backward swimming (knife) but also with forward motion (loc). Escape and rises are clustered together, as they are wrestling and head clamping. Chasing and head butts clustered with type 1 chirps. Type 3, jaw quivering and tail waving did not seem to cluster with any other behavior, probably due to their lower frequency within the dataset.

Overall, these results provided mutually reinforcing evidence indicating that chirps are produced more often during locomotion or scanning-related motor activity and confirm previous reports of a lower occurrence of chirping during more direct aggressive contact (as shown also by Triefenbach and Zakon, 2008; Hupé and Lewis, 2008). This mainly applies to type 2 chirps and backward swimming (knife; Triefenbach and Zakon, 2008), while larger chirps - often co-occurring (type 1 and 3 together, Figure 6C) - seem to be associated more frequently to fast-occurring behaviors (forward swimming and chasing), but not necessarily in their immediate proximity (as shown by the PSTH plots). Interestingly, despite this co-occurrence, type 3 chirps seem to cluster independently in the factor analysis, possibly suggesting differences in their usage pattern when compared to other large but faster chirps (type 1).

Earlier version: The significantly higher extent of chirping during swimming and locomotion, consistently confirmed by 4 different approaches (PSTH, TM, CN, MDS), suggests that chirping may be linked more strongly with scanning and environmental exploration than with a particular motivational state, thus confirming findings from our playback experiments.

Revised version: The significantly higher extent of chirping during swimming and locomotion, consistently confirmed by 4 different approaches (PSTH, TM, CN, MDS), suggests that - although chirp-behavior correlations may exist at time-scales larger than those here considered - chirping may be linked more strongly with scanning and environmental exploration than with a particular motivational state, thus confirming findings from our playback experiments.

Chirps significantly interfere with the beat and enhance electric image contrast

Earlier version: Given the stereotypical use of chirp types across various social contexts (Figure S2), the lack of chirp type-specific responses (Figure S3), their higher degree of auto-correlation of chirp time-series (Figure 5), the co-occurrence of chirps with locomotion and scanning (Figure 6), and the significant influence of beat frequency (i.e., DF) on chirp variance (Figure 2B, also emphasized by Walz et al., 2014), we propose that the function of chirps is more likely linked to physiological aspects of beat processing in the chirping fish rather than in the receiver. According to this hypothesis, the emission probability of different chirp types would depend on how detectable their interference with the beat is, specifically at the beat frequency at which they are produced (for an in-depth analysis see Walz et al., 2013, 2014). It is worth noting that the frequency-dependent effects of different chirps on the beat have already been extensively described (Walz et al., 2014; Petzold et al., 2016). To evaluate this possibility, we computed an estimate of the interference caused by chirps on beat periodicity (beat interference) using an array of Gaussian chirps across a wide range of beat frequencies (±300 Hz, Figure S8). Estimates of beat interference were made by calculating the ratio between the cumulative duration of the beat cycles affected by a given chirp (1 beat cycle corresponding to the interval between two consecutive beat peaks) over the cumulative duration of all the beat cycles within the time window used as a reference (700 ms, for other window sizes see Figure S9). This ratio was calculated for a wide range of chirp parameters: duration ranged between 10 and 400 ms, while FM ranged between 0 and 400 Hz (Figure S8A). Since chirps can be produced at any EOD phase (Walz et al., 2013), estimates of beat-interference were calculated for each chirp at 4 different phases (with 90° steps) and then averaged. As a result, we obtained a three-dimensional array of beat-interference values which could be matched to recorded chirps (N = 30486 chirps, from our first set of recordings) having the same parameters (duration, FM and DF; Figure S8B).

Revised version: Given the stereotypical use of chirp types across various social contexts (Figure S2), the lack of chirp type-specific responses (Figure S3), the higher degree of auto-correlation of chirp time-series (Figure 5), the co-occurrence of chirps with locomotion and scanning (Figure 6), and the significant influence of beat frequency (i.e., DF) on chirp variance (Figure 2B, also emphasized by Walz et al., 2014), we propose that the function of chirps is linked to physiological aspects of beat processing in the chirping fish. According to this hypothesis, the emission probability of different chirp types would depend on how detectable their interference with the beat is, specifically at the beat frequency at which they are produced (for an in-depth analysis see Walz et al., 2013, 2014). It is worth noting that the frequency-dependent effects of different chirps on the beat have already been extensively described (Walz et al., 2014; Petzold et al., 2016). To evaluate this possibility, we computed an estimate of the interference caused by chirps on beat periodicity (beat interference) using an array of Gaussian chirps across a wide range of beat frequencies (±300 Hz, Figure S8). Estimates of beat interference were made by calculating the ratio between the cumulative duration of the beat cycles affected by a given chirp (1 beat cycle corresponding to the interval between two consecutive beat peaks) over the cumulative duration of all the beat cycles within the time window used as a reference (700 ms, for other window sizes see Figure S9). This ratio was calculated for a wide range of chirp parameters: duration ranged between 10 and 400 ms, while FM ranged between 0 and 400 Hz (Figure S8A). Since chirps can be produced at any EOD phase (Walz et al., 2013), estimates of beat-interference were calculated for each chirp at 4 different phases (with 90° steps) and then averaged. As a result, we obtained a three-dimensional array of beat-interference values which could be matched to recorded chirps (N = 30486 chirps, from our first set of recordings) having the same parameters (duration, FM and DF; Figure S8B).

By plotting both beat-interference values and chirp abundance by DF, we found that emitted chirps were indeed often localized in correspondence with areas of high interference: at negative DFs, chirps were localized mainly in the type 2 range, whereas at positive DFs recorded chirps overlapped more often with the type 1 and 3 ranges (Figure S8C). This resembled what we already observed during the playback experiments in which EOD frequency ramps were used (Figure 4). Eventually, by calculating the normalized cumulative interference for each chirp type, a high degree of matching was found between the DF values at which the interference is higher and those at which chirps were actually recorded (Figure S8B).

Earlier version: A chirp type comparison (Figure S8D) further revealed that type 2 chirps at low or even negative DFs may induce similar interference levels as those induced by larger chirps at higher beat frequencies. This agrees with what previously reported using other parameters to estimate chirp-induced beat perturbations (i.e. chirp conspicuousness, Petzold et al., 2016) and could perhaps explain the lower occurrence of larger chirps in low frequency fish (more often in negative-DF encounters). In addition, areas of lower interference seem to be symmetrically localized around 60 Hz, the frequency level to which electroreceptors are most sensitive (Bastian, 1981) and at which chirp occurrence was lower (Figure S8B-D). However, this could be also due to the lower representation of these DFs values in this sample since responses to frequency ramps do not seem to convincingly validate this idea (Figure 4). Overall, these results corroborate the hypothesis that chirp production may be explained mainly by the effects chirps have on the beat perceived by the chirping fish.

Revised version: A chirp type comparison (Figure S8D) further revealed that type 2 chirps at low or even negative DFs may induce similar interference levels as those induced by larger chirps at higher beat frequencies. This agrees with what has been previously reported using other parameters to estimate chirp-induced beat perturbations (i.e. chirp conspicuousness, Petzold et al., 2016) and could perhaps explain the lower occurrence of larger chirps in low frequency fish (more often in negative-DF encounters). In addition, areas of lower interference seem to be symmetrically localized around 60 Hz, the frequency level to which electroreceptors are most sensitive (Bastian, 1981) and at which chirp occurrence was lower (Figure S8B-D). However, this could be also due to the lower representation of these DFs values in this sample since responses to frequency ramps do not seem to convincingly validate this idea (Figure 4). Overall, these results corroborate the hypothesis that chirp production may be explained mainly by the effects chirps have on the beat perceived by the chirping fish.

Differences in the beat interference elicited by different types of chirps were found: the interference induced by type 1 chirps is almost one order of magnitude higher than the estimates of the same parameter for type 2 chirps, but one order of magnitude lower than the beat interference induced by type 3 chirps. Rises induce interferences in the range of type 1 chirps (not shown).

Earlier version: Obviously, measuring chirp-triggered beat interferences by using an elementary outlier detection algorithm on the distribution of beat cycles does not reflect any physiological process carried out by the electrosensory system and can be therefore used only as an oversimplified estimate. A more realistic representation of beat processing and chirp-triggered beat perturbations can be done though electric images (EIs), which are 2D representations of the current flowing through electroreceptors located on the fish skin at any given moment in time (Kelly et al., 2008). In the context of social interactions, EIs represent the projection of the electric field generated by two interacting fish onto the array of electrosensory receptors, at the net of the environmental perturbations caused by objects, motion and other sources of noise (Kelly et al., 2008; Fotowat et al., 2013).

Revised version: Obviously, measuring chirp-triggered beat interferences by using an elementary outlier detection algorithm on the distribution of beat cycles does not reflect any physiological process carried out by the electrosensory system and can be therefore used only as an oversimplified estimate. A more realistic representation of beat processing and chirp-triggered beat perturbations can be done through electric images (EIs), which are 2D representations of the current flowing through electroreceptors located on the fish skin at any given moment in time (Kelly et al., 2008). In the context of social interactions, EIs represent the projection of the electric field generated by two interacting fish onto the array of electrosensory receptors, at the net of the environmental perturbations caused by objects, motion and other sources of noise (Kelly et al., 2008; Fotowat et al., 2013).

Both the heterogeneous distribution of electroreceptors as well as the asymmetric geometry of the electric field contribute to the further alteration of electric images during locomotion (Rasnow, 1996; Yu et al., 2012). While the EOD is one order of magnitude stronger at the tail end compared to the head, the EOD harmonic components also change along the body axes. This means that the perception of both beat contrast and beat frequency can change at different body locations (Rasnow et al., 1996). This explains also why electric images are easily altered by motion and why, in the presence of anisotropic electric field diffusion caused by soluble materials or by environmental objects, beat perception is not an homogeneous process throughout the whole population of electroreceptors. In fact, depending on the fish relative positioning and orientation, electric images can be subject to significant degradation or near cancellation, contingent on the phase of interacting EODs (Kelly et al., 2008). To prevent this, image contrast can, in principle, be enhanced either by adjustments of the body position or through other mechanisms amplifying the responses of either electroreceptor afferents (P-units) or downstream circuits (ELL pyramidal cells). Since the beat AMs generated by the chirps always trigger reliable responses in primary electrosensory circuits (pyramidal cells in the ELL respond to both increases and decreases in beat AM), any chirp-triggered AM causing a sudden change in P-unit firing could potentially amplify the downstream signal (Marsat and Maler, 2010) and thus enhance EI contrast.

Earlier version: To simulate the impact of chirps on EIs and assess their influence on image contrast, we computed EIs using actual chirp recordings from paired fish interactions (N = 4 pairs, 297 chirps). These were then compared with beats having identical AM and frequency parameters (Figure 7A,B), resulting in a pair of EIs for each chirp type. Comparisons of chirp-EI with beat-EI without chirps consistently showed that chirps led to increased contrast in electric images (Figure 7C). Importantly, this effect is mainly evident for sender fish, if compared to receivers (Figure 7D). This suggests that chirp production could be potentially advantageous for electrosensory processing in scenarios in which beat amplitude is either low or poorly resolved. Overall, these data further support the idea of chirping being used for beat processing purposes. Both the estimates of chirp-induced beat interference as well as the more realistic representations of chirp-triggered electric images confirm that chirping exerts conspicuous and chirp type-specific effects on the beat, easily detectable by electrosensory afferents and therefore usable to complement beat-processing.

Revised version: To simulate the impact of chirps on EIs and assess their influence on image contrast, we computed EIs using actual chirp recordings from paired fish interactions (N = 4 pairs, 297 chirps). These were then compared with beats having identical AM and frequency parameters (Figure 7A,B), resulting in a pair of EIs for each chirp type. Comparisons of chirp-EI with beat-EI without chirps consistently showed that chirps led to increased contrast in electric images (Figure 7C). Importantly, this effect is mainly evident for sender fish and much less so for receivers (Figure 7D). This suggests that chirp production could be potentially advantageous for electrosensory processing in scenarios in which beat amplitude is either low or poorly resolved. Overall, these data further support the idea of chirping being used for beat processing purposes. Both the estimates of chirp-induced beat interference as well as the more realistic representations of chirp-triggered electric images confirm that chirping exerts conspicuous and chirp type-specific effects on the beat, easily detectable by electrosensory afferents and therefore usable to complement beat-processing.

Chirps enhance the contrast of electric images in the sender but not in the receiver fish.

Earlier version: A: Heatmaps representing the electric field (V/cm) generated by a sender fish alone (left) and by two interacting fish (right, fish length = 15 cm, distance = 10 cm). The electric field lines induced by the sender fish’s own EOD (gray) and that of the receiver conspecific (white) are shown in black. B: Normalized heatmaps superimposed on the sender fish contour representing the boundary element model (BEM) simulation of the electric image as a result of the field interactions measured on the sender fish skin. The electric images are calculated by subtracting the transcutaneous current in the presence and absence of the receiver fish. C: Sum of the electric images modeled on the skin of the sender fish and generated by different chirp types produced by the same subject (N = 297 chirps). These effects are compared to the respective beats without chirps (beat range ca. 40-300 Hz). D: Net effect of chirps on electric images represented as absolute current changes induced by chirps and compared to their carrier beats alone. Data for sender (top) and receiver fish (bottom) are displayed separately. Note that chirps do not elicit significant changes in the electric images calculated for the receiver fish.

Revised version: A: Heatmaps representing the electric field (V/cm) generated by a sender fish alone (left) and by two interacting fish (right, fish length = 15 cm, distance = 10 cm). The electric field lines induced by the sender fish’s own EOD (gray) and that of the receiver conspecific (white) are shown in black. B: Normalized heatmaps superimposed on the sender fish contour from the same instant as in A representing the boundary element model (BEM) simulation of the electric image as a result of the field interactions measured on the sender fish skin. The electric images are calculated by subtracting the transcutaneous current in the presence and absence of the receiver fish. C: Electric image changes calculated over a time window of 500 ms. The ΔEI represents the cumulative change in current flow across the skin in presence of a conspecific minus the current flow measured in its absence. The electric image of the conspecific (body and beat only) is shown on the top left. The other images refer to different types of chirps being produced on the background of the same beat by the same subject (N = 297 chirps; beat range ca. 40-300 Hz, both fish of the pairs chirped). D: Net effect of chirps on electric images represented as the sum of current changes induced by chirps on the whole body (800 nodes) and compared to their carrier beats alone. Data for sender (top) and receiver fish (bottom) are displayed separately. Note that chirps do not elicit significant changes in the electric images calculated for the receiver fish.

Chirping is significantly affected by locomotion and social distance

To obtain more insights into the potential role of chirps in beat processing we evaluated chirping and different aspects of swimming activity using a novel environment exploration assay. We assessed chirp production by brown ghosts freely swimming while localizing a beat source (a caged conspecific). This allowed to evaluate aspects of fish locomotion and social behavior neglected in previous experiments, potentially explaining high inter-individual chirping variability (e.g., thigmotaxis, novelty-seeking, sociability). If chirping relates to geometrical aspects of beat processing (distance, orientation, direction, and motion), a bias for any of these should be evident in these experiments. If chirps enhance beat processing and conspecific electrolocation, for instance, chirping should occur within beat detection range but at a certain distance. Alternatively, if chirps help resolve other parameters of electric field geometry, they may be produced at specific angles or when fish are in close proximity.

Fish were allowed to swim freely (for 30 min) in a 160 x 80 cm arena containing shelters (plastic tubes and plants). The caged conspecific was located in a mesh tube on one side of the arena (Figure 8A). A “novel object” (a 1 x 5 cm graphite rod) was hidden behind a barrier and used to test novelty seeking. Shelters consisted of transparent PVC tubes and plastic plants. Different regions of interest (ROIs) were used to score the occurrence of fish swimming behavior around the shelter area, the tank walls (within 5 cm distance), the novel object and the conspecific (Figure 8A). The time spent swimming in the arena’s open space (5 cm away from the walls) was also evaluated. Behavioral observations were conducted considering the subject’s sex, as it is the behavioral variable most directly correlated to DF and beat frequency.

Chirping during novel environment exploration.

A: Diagram of the recording arena showing the criteria used to define the different regions of interest (ROIs): 1) the presence of shelters (PVC tubes), 2) the proximity to the tank walls (distance < 5 cm), 3) the presence of a fish (caged in a mesh tube), 4) the presence of an “unknown/novel” conductive object (a 3 cm piece of graphite). B: Proportion of time spent in the different ROIs (N = 14 females, Friedman X2 22,9 p < 0.001 wall vs open p < 0.001, social vs object p = 0.02, open vs shelter/social/object < 0.01; N = 15 males, Friedman X2 19,8 p < 0.001 wall vs open p = 0.049, wall vs shelter/object p < 0.05, open vs shelter/social/object p < 0.001). C: Chirp locations (red) overlaid to the heatmaps showing the average swimming activity of females and males (for the specific locations of different types of chirps see Figure S10). D: Polar histograms showing the angles between the two fish during chirping. Angles are referred to the X-axis and are sorted based on chirp type and sex (male = blue, female = red). E: Histograms of the chirping distances in males (blue) and females (red) relative to different types of chirps.

We found that on average brown ghosts spent little time exploring the caged conspecific (12% of the time) and the shelter (14%), while slightly more attention was paid to the novel object (18%), swimming along the walls (ca. 20%) or exploring other objects in the “open” environment (ca. 35% of the time). In particular, female brown ghosts spent more time swimming in the open space areas of the arena, as opposed to swimming closer to the tank walls (thigmotaxis; Figure 8B). Female subjects also spent more time investigating the novel object than the conspecific (Figure 8B). Male behavior showed a different trend as males spent comparatively more time swimming close to the tank walls and more time in contact with conspecifics (Figure 8B). Possibly as a direct consequence of this, they also produced a higher number of chirps, compared to their female counterparts. Most chirps were produced in close proximity to the caged conspecific (Figure 8C). Analysis of the angles and distances of fish during chirps revealed that first, chirps were mostly produced when subjects were located at 30°, 150° or 330° angles relative to the axis defined by the mesh tube (Figure 8D), and second, most chirps were produced at a distance of less than 25-50 cm (Figure 8E). While the angle bias could be due to the circular loops described by fish swimming around the conspecific, the rather limited distance range at which chirps are produced could be explained by the range of the EOD field (beat amplitude reaches 1% at around 30 cm; Fotowat et al., 2013; Henninger et al. 2018; Benda, 2020; see appendix 1) and confirms previous observations (Hupé and Lewis, 2008; Zupanc et al., 2006; Henninger et al., 2018).

Interestingly, chirps of different types were not all produced at similar angles: while type 1 and type 2 chirps seem to be more often produced at 30-150 angles, type 3 chirps - produced in this experiment only by males - seem to be used in the ±30° range (i.e. on the right side of both the arena and the plane representing the angles). Rises - in both males and females - seemed more often to be produced at multiples of 30° located on the left side of the cartesian plane (30°, 150° and 210°). All chirp types, except for rises, were produced within a 50 cm radius from the conspecific fish (Figure 8E). Rises were produced almost ubiquitously, by both sexes, and most of the time at locations far away from the caged fish (Figure S10).

Factor analysis (FAMD) confirmed the already known sex-dimorphisms in chirp distribution (males chirp more than females and more often towards other males, Figure S11). Collectively, caged fish produced more chirps than freely swimming subjects (4410 vs. 3483 chirps). In addition, variables associated to beat frequency (such as DF, EOD frequency, sex, chirp type) explain a good percentage of chirp variance (73,7% of variance explained as opposed to the 55% explained including all other variables 55%; see also Figure 2B). A relatively high loading on principal components was also found for variables associated with locomotion, such as distance and time spent exploring different tank areas (Figure S11B). Accordingly, chirp clustering is affected by the EOD frequency difference of the emitting fish (DF) and by distance (Figure S11C, insets).

Earlier version: Modelling of field current and iso-potential lines using dipole models of our fish allows to visualize more intuitively how both chirp angles and distances can be explained by the electric field geometry (Figure S12A). Most chirps (90%) are in fact produced within a distance corresponding to 1% of the maximum field intensity (i.e. roughly 30 cm; Figure S12B), indicating that chirping occurs way above the threshold distance for beat detection (i.e. roughly in the range of 60-120 cm, depending on the study; see appendix 1: Detecting beats at a distance). This effect is even more evident when chirp locations are compared to the values of field intensity measured at different distances (Figure S12C,D). By estimating the interference each chirp type would induce on the beat, based on its decay law (Figure S12D), one can easily see how the two are linearly and directly correlated (Figure S12E-H). Overall, given that most chirps are reliably produced within very short range, they may not be useful for improving beat detection. Conversely, the orientation bias for chirp locations suggests that chirps are used during specific swimming trajectories and thus likely occurring while brown ghosts are trying to resolve other spatial parameters, such as the orientation, direction and possible motion of the beat source. In all this, rises may represent an exception as their locations are spread over larger distances and even in presence of obstacles potentially occluding the beat source (i.e. shelters, plants, or walls), all conditions in which beat detection or beat processing could be more difficult or simply interrupted (an idea coherent with the production of rises right at the end of EOD playbacks; Figure S5). Taken together, these observations may possibly explain why female subjects, apparently less motivated in exploring social stimuli, also produce considerably fewer chirps unless they are forced into social interactions (Figure S1E-F). More broadly, these findings indicate that chirping, besides being strongly correlated to the frequency of the beat generated by the interacting fish, is also significantly affected by locomotion and by spatial relationships between subjects within beat detection range. This further corroborates the hypothesized role of chirps in beat processing.

Revised version: Modelling of field current and iso-potential lines using dipole models of our fish allows to visualize more intuitively how both chirp angles and distances can be explained by the electric field geometry (Figure S12A). Most chirps (90%) are in fact produced within a distance corresponding to 1% of the maximum field intensity (i.e. roughly 30 cm; Figure S12B), indicating that chirping occurs way above the threshold distance for beat detection (i.e. roughly in the range of 60-120 cm, depending on the study; see appendix 1: Detecting beats at a distance). This effect is even more evident when chirp locations are compared to the values of field intensity measured at different distances (Figure S12C,D). By estimating the interference each chirp type would induce on the beat, based on its decay law (Figure S12D), one can easily see how the two are linearly and directly correlated (Figure S12E-H).

Added in the revised version: Overall, given that most chirps are reliably produced within very short range, they may not be useful for improving beat detection. Conversely, the orientation bias for chirp locations suggests that chirps are used during specific swimming trajectories and thus likely occurring while brown ghosts are trying to resolve other spatial parameters, such as the orientation, direction and possible motion of the beat source. In all this, rises may represent an exception as their locations are spread over larger distances and even in presence of obstacles potentially occluding the beat source (i.e. shelters, plants, or walls), all conditions in which beat detection or beat processing could be more difficult or simply interrupted (an idea coherent with the production of rises right at the end of EOD playbacks; Figure S5). Taken together, these observations may possibly explain why female subjects, apparently less motivated in exploring social stimuli, also produce considerably fewer chirps unless they are forced into social interactions (Figure S1E-F). More broadly, these findings indicate that chirping, besides being strongly correlated to the frequency of the beat generated by the interacting fish, is also significantly affected by locomotion and by spatial relationships between subjects within beat detection range. This further corroborates the hypothesized role of chirps in beat processing.

Chirping increases with environmental complexity

We hypothesized that if chirps are used as probes to refine beat processing and/or the resolution of electric images produced by conspecifics, higher levels of environmental complexity would likely result in higher chirping rates (Siemers et al., 2009). To test this idea, we used both playback experiments and recordings of interacting fish pairs. The objective was to evaluate the chirping responses of brown ghosts to EODs in presence of environmental clutter which would distort electric fields in tangible but unpredictable ways (von der Emde, 2006).

We first recorded social interactions using the same divided tank configuration as in our previous recordings. Fish pairs (N = 6) were recorded in 3 different conditions: 1) with lights turned ON and no object in the tank other than the mesh divider placed between the 2 fish (lights ON - clear), 2) with lights switched OFF (lights OFF - clear) and 3) with lights OFF and the addition of shelter PVC tubes and plastic plants filling almost entirely the two tank compartments (lights OFF – cluttered; Figure 9A). Fish pairs were recorded for 15 minutes with no previous knowledge of either the environment or their social partner. Trials were shuffled in a randomized order and were spaced by 5-minute inter-trial intervals. Chirp counts were evaluated at the end of the trials. As expected, the total chirp rate increased in absence of light (Zupanc et al., 2001). More strikingly, the increase was even more pronounced when clutter was present in the environment (Figure 9B). In all cases, type 2 chirps were the most often produced type (Figure 9C). A slight (non-significant) increase in type 2 and type 3 chirps could be observed under cluttered conditions. These results clearly indicate that chirping is affected by environmental complexity, since in these recordings aggression- or courtship-related behaviors were either absent or randomly occurring in the three conditions and both fish had equal environmental and social experience.

Effect of environmental clutter on interacting fish pairs.

A: Recording of fish pairs (N = 6) in environments of different sensory complexity: lights ON = clear tank environment and direct illumination, lights OFF = clear tank environment, no illumination, lights OFF + clutter = no illumination and cluttered environment. B: Total amount of chirps produced in each condition and normalized on the lights ON session (green bars; Friedman’s p = 0.053; lights ON vs lights OFF + clutter = 0.024, lights OFF vs lights OFF + clutter = 0.051). Chirp counts relative to trials sorted in chronological order (1-3) are shown in gray. C: The box plots show the normalized chirp type counts relative to each session (lights ON, Friedman’s X2 = 21.9 p < 0.001, pairwise comparisons 1 vs 2 p = 0.006, type 1 vs 3, type 1 vs 4, type 2 vs 3, type 2 vs 4 p < 0.001; lights OFF Friedman’s X2 = 21.8 p < 0.001 pairwise comparisons type 1 vs 2, type 1 vs 3, type 1 vs 4, type 2 vs 3, type 2 vs 4 p < 0.001; clutter Friedman’s X2 = 19.5 p < 0.001 pairwise comparisons type 1 vs 2, type 1 vs 4, type 2 vs 3, type 2 vs 4 p < 0.001). D: Results of playback experiments in which EOD mimics were either directly detectable through a fine mesh barrier (clear) or more indirectly due to a barrier of plastic plants interposed between the mesh and the EOD source (cluttered). Clear and cluttered trials were presented in random succession (N fish = 6, 10 trials each, 60 sec ITI). E: Total chirp counts in the 2 conditions are normalized on the total amount of chirps produced by each subject (Wilcoxon, p = 0.025). F: Boxplots showing the chirp type composition of each condition (clear Friedman’s X2 = 17.4 p < 0.001 pairwise comparisons type 1 vs 3 p = 0.048, type 1 vs 4, type 2 vs 3, type 2 vs 4 p < 0.001; cluttered Friedman’s X2 = 29.9 p < 0.001 pairwise comparisons type 1 vs 3, type 1 vs 4, type 2 vs 3, type 2 vs 4 p < 0.001; Wilcoxon type 1 clear vs type 1 clutter p = 0.034). Tank sizes: A-C: 30 x 80 cm; D-F: 160 x 50 cm.

Nonetheless, even in conditions of low “electrical detectability”, fish can sense each other using visual cues, olfactory and mechanical cues induced by water movements during swimming. Therefore, chirp production in these experiments could also be affected by a range of sensory inputs co-occurring in an unnaturally confined space and for the whole duration of the recordings (15 min). To rule out these possibilities, we conducted playback experiments in a larger enclosure (160 x 50 cm), allowing fish to detect playback EODs from a wider range of distances. Playback electrodes were not directly accessible but placed behind a mesh barrier located at either side of the enclosure (Figure 9D). Environmental clutter (plastic plants) was interposed between the electrodes and the mesh divider at a random location in each trial. We found that chirp rates increased in response to playback signals specifically when they were partially concealed by clutter Figure 9E). Notably, this effect seems to be more evident for larger chirps (type 1 and type 3; Figure 9F) although combined effects of chirp types and clutter have not been found. Overall, these results indicate that chirping is significantly affected by the presence of environmental clutter partially disrupting - or simply obstructing - the processing of beat related information during locomotion. They further confirm electric signals alone are sufficient to trigger the chirping increase in cluttered conditions. All chirp types seem to be equally affected, with type 2 chirps being the most common type produced also in cluttered conditions.

The increased utilization of chirping during locomotion around partially concealed EOD sources mirrors the higher probing rates typically observed in complex environments (Obrist, 1995; Schnitzler and Kalko, 2001; Slabbekoorn and Peet, 2003; Siemers et al., 2009). This supports the notion of chirps as self-referenced probing cues, potentially employed to optimize short-range aspects of conspecific electrolocation, such as object size, orientation, and direction - a hypothesis that will certainly be explored in future studies.

Discussion

In this study, we examined the factors influencing the production of EOD frequency modulations (chirps) in brown ghost knifefish. We gathered a substantial dataset of chirps recorded during various social interactions, evaluating the impact of territory ownership, social hierarchy, sex, reproductive state, and social experiences. Analyzing this data, we found that the most influential factors in chirping are primarily beat frequency (DF) and experience (both environmental and social). Chirp patterns mainly consist of repetitions of the same chirp type, heavily influenced by DF. Responses to playback chirps inserted into a sinewave stimulus of a particular DF are indistinguishable in terms of chirps emitted by the fish and exploratory behavior (approach/avoidance). Our behavioral observations of freely swimming fish indicate that chirping typically occurs at short ranges, often coinciding with active social approaches. Building upon these findings, we suggest that chirps serve as self-directed signals (sensory probes), aiding in the processing of spatial information derived from beats.

Existing functional hypotheses

Since the earliest reports (Larimer and MacDonald, 1968), chirps have been assigned behavioral meaning as communication signals based on the following observations (mostly in brown ghosts): a) fish chirp mainly during social encounters; b) males chirp considerably more than females (Dye, 1987; Bastian et al. 2001, Tallarovic and Zakon 2002, Dunlap and Larkins-Ford, 2003; Kolodziejski et al. 2007) and are also considerably more aggressive; c) during male interactions (often agonistic) smaller chirps are produced (i.e. shorter duration and lower frequency excursion: type 2) whereas larger chirps (types 1 and 3) are used more often during female-male encounters (presumably courtship-driven); d) high chirp rates are observed in mating pairs, too, with a relatively higher production of large chirps (i.e. type 1 and type 3; Hagedorn and Heiligenberg, 1985; Henninger et al., 2018). This led to a dualistic semantic subdivision of chirps, with small chirps implied in agonistic interactions and larger chirps pertaining to different aspects of courtship and mating.

Earlier version: Although intuitive, this idea raises the following concerns: first, several factors (sex and the magnitude of beat-interference) covary with the main corollary of chirping, i.e. beat frequency. These covariates are not easy to disentangle, which makes it difficult to separate neuroendocrine or other behavioral factors from the biophysical features of EOD fields affecting detection and localization of conspecific fish. Second, observations in a few species are generalized to all other gymnotiforms without testing whether chirping may have similar functions in other species (Turner et al., 2007; Smith et al., 2013; Petzold et al., 2016). Third, social encounters do not only involve communication but also higher levels of locomotion and reciprocal exploration, which could also be linked to chirping activity, but have not been considered in the past. Last, and most importantly, no causal evidence in support of any of the hypothesized functions of the different types of chirps has ever been provided.

Revised version: Although intuitive, this idea raises the following concerns: first, several factors (sex and the magnitude of beat-interference) covary with the main corollary of chirping, i.e. beat frequency. These covariates are not easy to disentangle, which makes it difficult to separate neuroendocrine or other behavioral factors from the biophysical features of EOD fields affecting detection and localization of conspecific fish. Second, observations in a few species are generalized to all other gymnotiforms without testing whether chirping may have similar functions in other species (Turner et al., 2007; Smith et al., 2013; Petzold et al., 2016). Third, social encounters do not only involve communication but also higher levels of locomotion and reciprocal exploration, which could also be linked to chirping activity, but have not been considered in the past. Fourth, studies supporting a dual function of chirping (i.e. aggression and courtship), only rarely addressed this hypothesis using breeding animals. Last, and most importantly, no causal evidence in support of any of the hypothesized functions of the different types of chirps has ever been provided.

The idea of chirps being used as probes, as proposed here, is not entirely new: due to the tight correlation between chirp features and beat frequency (mainly chirp frequency modulation), Walz and coauthors previously hypothesized a role of chirps in determining the sign of the beat frequency (Walz et al., 2014). This indeed may represent an ambiguous parameter as two opposite frequency differences (or DFs) result in the same beat frequency: e.g., consider a fish producing an EOD of 750 Hz (EOD1) and interacting with conspecifics with an EOD of either 700 Hz (EOD2) or 800 Hz (EOD3). The two beats could be indistinguishable (assuming - perhaps not realistically - that the same mechanism involved in DF discrimination at lower DF values would not work in this case; Metzner, 1999). However, since any frequency modulation of EOD1 will have different effects on the beat frequency relative to EOD2 and EOD3, chirps could indeed solve the problem of disambiguating the sign of DF (which in brown ghosts could mean no less than discriminating between male and female conspecifics). However, it is also possible to obtain beat frequency information more reliably with a more gradual and slow change in the fish’s own EOD frequency (which knifefish are capable of producing, see for instance Rose and Canfield, 1991). More banally, the average chirp rates do not seem to realistically reflect such an aim (at low beat frequencies, it is not rare for a fish to chirp more than 100 times per minute).

More recent studies have further suggested that patterns of specific chirp types could be a salient signal used for communication purposes in the context of mating (Henninger et al., 2018). Yet, results obtained from our recordings do not support the existence of a chirp-based syntax in social interactions in which breeding conditions were simulated: fish pairs acclimated to breeding conditions (but not mating) used non-patterned chirp repertoires similar to fish pairs recorded in non-breeding conditions. One exception to the monotony of the chirping repertoires observed in this study is given by male-female pairs: in this case, type 1 and type 3 chirps are used more often, while the relative amount of type 2 chirps is lower (Figures 5, S2C,D). However, this effect is likely to be a result of the combined effect of the pronounced dependency of chirp type on DF together with the higher chirp rates observed in males (i.e. if chirp variety increases with DFs and males chirp more, chirp type variety will be higher when males interact with high beat frequencies). Moreover, the somewhat stereotypical transition between large and small chirps that can be observed during repeated exposure to playback frequency ramps, appears a convincing argument against the communication-hypothesis.

Inconsistencies between behavior and hypothesized signal meaning

In our playback experiments we observed no differences in the responses of fish of the two sexes to playback stimuli, regardless of whether they contained a particular chirp type or not: test subjects chirped at comparable rates in response to all types of EOD mimics (containing type 1 chirps, type 2 chirps or rises), confirming that the main factor determining type choice is DF and therefore self-referenced sensory requirements (as opposed to different semantic contents). This idea is further corroborated by the known symmetric distribution of chirp types within the ±400 Hz DF range (Bastian et al., 2001; Engler and Zupanc, 2001; Triefenbach and Zakon, 2008; Kolodziejski et al., 2007). Nonetheless, it’s plausible that playback stimuli, as employed in our study and others, may not faithfully replicate natural signals, thus potentially influencing the reliability of the observed behaviors. However, EOD waveform and harmonic components – which make natural EODs different from playback mimics - do not seem to play a role in EOD discrimination as EOD frequency alone does (Dunlap and Larkins-Ford, 2003; Fugère and Krahe, 2010). Moreover, artificial EODs are widely used in electrophysiological studies as they elicit comparable responses to natural stimuli in the electrosensory system (e.g., Benda et al., 2006, Marsat et al., 2009). Future studies might consider replicating these findings using either natural signals or improved mimics, which could include harmonic components (excluded in this study).

Based on our beat interference estimates, we propose that the occurrence of the different types of chirps at more positive DFs (such as in male-to-female chirping) may be explained by their different effect on the beat (Figure S8; Benda et al., 2006; Walz et al., 2013).

Earlier version: In line with our behavioral data, electrophysiological recordings conducted at all main nodes of the electrosensory pathway did not show consistent chirp-type specific responses in either peripheral or central brain areas (Metzen and Chacron, 2017; Allen and Marsat, 2018; Metzen et al., 2020). Instead, the temporal correlations between chirps and other events (locomotion-related) seem to be more salient factors, as recently proposed for midbrain circuits whose phasic responses to moving objects are used to timestamp their occurrence during spatial navigation (Wallach et al., 2018). Recordings from interacting fish pairs confirmed the absence of any significant correlation between chirp type choice and behavioral context (Figure S2) although the variance of chirp parameters appears to be significantly affected by this factor (Figure 2). This may suggest that the effect of behavioral context is mainly detectable in the number of chirps produced (Figure S1), rather than the type (Figure S2).

Revised version: In line with our behavioral data, electrophysiological recordings conducted at all main nodes of the electrosensory pathway did not show consistent chirp-type specific responses in either peripheral or central brain areas (Metzen and Chacron, 2017; Allen and Marsat, 2018; Metzen et al., 2020). Instead, the temporal correlations between chirps and other events (locomotion-related) seem to be more salient factors, as recently proposed for midbrain circuits whose phasic responses to moving objects are used to timestamp their occurrence during spatial navigation (Wallach et al., 2018). Recordings from interacting fish pairs confirmed the absence of any significant correlation between chirp type choice and behavioral context, except for those cases characterized by higher beat frequencies (Figure S2). This suggests that the effect of behavioral context highlighted in our factor analysis (Figure 2) is mainly due to the number of chirps produced (Figure S1), rather than their type (Figure S2).

Type 1 and type 3 chirps (i.e. large chirps), have been often suggested to be courtship related signals because they are emitted predominantly by opposite sex pairs (i.e. fish pairs characterized by a higher DF value; Hagedorn and Heiligenberg, 1985; Henninger et al., 2018). However, the symmetric distribution of their production probability around the lowest DF values (relative to the sender EOD frequency), suggests that this conclusion may not be always valid. In addition, here we show not only that large chirps are also consistently produced at very low DFs but also that the relative proportion of different chirp types does not seem to differ in social pairings implying potentially very different internal states in the interacting fish (breeding, aggression, territorial competition, etc.). Both results are further weakening the link between large chirps and courtship displays.

Type 2 chirps are also considered to be significantly correlated with aggressive encounters, although it is not clear whether they represent purely aggressive, possibly deterring, signals (Zupanc, 2002; Triefenbach and Zakon, 2008) or attempts to de-escalate aggression and communicate submission (Hupé and Lewis, 2008). Our results do not support either of these hypotheses: type 2 chirps are produced at higher rates at low beat frequencies (a feature of any encounter in which fish with similar EOD frequency approach each other and which has nothing to do with their intentions), they are more often produced by newly introduced fish (which are not necessarily more aggressive but could just be more explorative), less in experienced fish (which also move less; Figures S1A,D, S2), more by males when separated from females (males are more motivated to interact socially; Figure 8) and more in females when freely swimming (females are more motivated to escape when exposed to males directly).

Earlier version: Notably, rises reliably occurred at the end of the playback stimulations (as also reported by Kolodziejski et al., 2007), which could be explained by their use as probes to assess signal presence in case of abrupt interruptions or weakening of EOD mimics (either caused by objects occluding temporarily the EOD field - clutter - or by playback termination). Interestingly, rises were considered by other authors as “proximity signals”, although most likely with a different connotation (Hupé & Lewis 2008). Rises larger than those here considered (perhaps more similar to the “gradual frequency rises” reported by Zupanc et al., 2001) were recently associated with chasing attempts (15% of the times) and physical contact (less than 5% of the times; Raab T, personal communication). However, these figures do not seem to be consistent enough to allow reliable conclusions about the behavioral meaning possibly associated to these frequency modulations. Previous experiments reported the use of long rises (larger events compared to those we have labeled as rises) during agonistic interactions (Hupé & Lewis, 2008; Raab et al., 2021). However, it is not clear whether these are aggressive signals (as proposed by Hupé and Lewis, 2008, Raab et al., 2021) or de-escalating submissive signals (as proposed by Hopkins, 1974, and Serrano-Fernández, 2003). In our recordings rises are produced more or less with equal probability by female and male subjects (42.5% of all rises are produced by males, 57.5% by females) and yet, compared to other chirp types, their relative probability is higher in females (1.2% of all chirps in females vs 0.12% in males), although female brown ghosts are less frequently engaged in aggressive interactions.

Revised version: Notably, rises reliably occurred at the end of the playback stimulations (as also reported by Kolodziejski et al., 2007), which could be explained by their use as probes to assess signal presence in case of abrupt interruptions or weakening of EOD mimics (either caused by objects occluding temporarily the EOD field - clutter - or by playback termination). Interestingly, rises were considered by other authors as “proximity signals”, although most likely with a different connotation (Hupé & Lewis 2008). Rises larger than those here considered (perhaps more similar to the “gradual frequency rises” reported by Zupanc et al., 2001) were recently associated with chasing attempts (15% of the times) and physical contact (less than 5% of the times; Raab T, personal communication). However, these figures appear too small to allow reliable conclusions about the meaning potentially associated with these frequency modulations. Previous experiments reported the use of long rises (larger events compared to those we have labeled as rises) during agonistic interactions (Hupé & Lewis, 2008; Raab et al., 2021). However, it is not clear whether these are aggressive signals (as proposed by Hupé and Lewis, 2008, Raab et al., 2021) or de-escalating submissive signals (as proposed by Hopkins, 1974, and Serrano-Fernández, 2003). In our recordings rises are produced more or less with equal probability by female and male subjects (42.5% of all rises are produced by males, 57.5% by females) and yet, compared to other chirp types, their relative probability is higher in females (1.2% of all chirps in females vs 0.12% in males), although female brown ghosts are less frequently engaged in aggressive interactions.

Taken together, these considerations strongly indicate that the primary determinants of chirp rates are the locomotor activity levels and the necessity to locate other fish through the beat. In our experiments, the choice of chirp type appeared largely unaffected by the behavioral context in which chirps were recorded, while the degree of social interaction, as observed in comparisons between novel and experienced fish pairs, played a bigger role.

Probing with chirps

In this study we explore the possibility that chirps could improve the processing of spatial parameters associated to conspecific localization by enhancing beat processing. The frequent use of chirping during close physical interactions, suggests that whatever role chirps may play in beat processing, this may not necessarily be related to the mere localization of a beat source at a distance (such as most echolocation calls) but perhaps to other features detectable at close-range (such as size, orientation, direction and motion). This could be achieved in at least 3 different ways. First, chirping could temporarily adjust beat frequency to levels better detectable by the electrosensory system (Bastian et al., 2001). Second, chirps could be used to improve active electrolocation by briefly enhancing beat temporal resolution (all types of chirps will increase ELL pyramidal cell firing rates, acting on different input lines; Marsat and Maler, 2010). Third, as previously discussed within this study, chirps could be used to enhance electrosensory responses to the beat. This may be achieved by simply enhancing the output of the primary electrosensory afferents (through synchronization of different types of input; Benda et al., 2006) or by phase-shifting the carrier EOD in correspondence of the destructive components of the beat cycles (i.e. the beat minima). Because the phase of interacting EODs is affected by both position and movement, phase shifts induced by chirps could be also aimed at recovering blind-spots forming on the electric images as a consequence of particular body positions, obstacles or even by motion (we show in fact that chirps can improve the EI contrast). Although chirp production has been previously reported to be phase-invariant (Metzen and Chacron, 2017), invariance here means that chirps are not produced at a specific EOD phase. Yet, this does not exclude the possibility that chirps could be used to briefly shift the EOD phase in order to avoid disruptive interferences caused by phase opposition (at the level of p-units). Therefore, this chirps-as-saccades idea could be still a valid hypothesis if one considers the effects of beat phase adjustment only at specific body locations (such as the head, which can be considered the “electrosensory fovea”; Carr et al., 1982; Castelló et al., 2000). In this view, chirping could be seen as an attempt to adaptively modify the incoming sensory input and maintain optimal contrast levels or a coherent sensory flow, as sniffing does for olfactory perception (Hahn et al., 1994), whisking and finger movement for tactile sensing (Krammer et al., 2020) and saccades for vision (Najemnik and Geisler, 2005). Future studies focused on the details of the EOD geometries and fish spatial coordinates during chirping will shed more light onto these matters.

In theory, chirps could also be used to improve electrolocation of objects as well (as opposed to the processing of the beat). Compared to other electric fish using pulse-type EODs, the frequency content of wave-type EODs is narrowly focused on limited components. The larger signal band-width (i.e. frequency content) of pulse EODs affords better object discrimination, when compared to wave-type EODs. However, given their lower production rates compared to the frequency of wave-type EOD cycles, EOD pulses grant a lower sampling rate and thus a poorer temporal resolution (Crampton, 2019; Bastian, 1976; Watson & Bastian, 1979). To compensate for this, EOD pulses are emitted at higher frequency during electrolocation bouts and locomotion in general (Jun et al., 2014). Conversely, for wave-type fish, chirping could represent a strategy to temporarily compensate for the lower frequency resolution while still being able to resolve EOD perturbations with a good temporal definition. Conducting combined video and playback experiments with a more specific focus on target-reaching behaviors could potentially elucidate whether this is the case. In addition, testing beat detection capabilities at different frequencies (or DFs) would be another important experimental approach to evaluate the probing hypothesis.

Earlier version: If this idea finds further support, the question arises as to why not all gymnotiforms electrolocate using the same strategy. It may be that a trade-off exists between space- and time-resolution in the evolution and maintenance of electrosensory systems (Crampton, 2019). While broadband pulse signals may be useful to capture highly complex environments rich in foliage, roots and vegetation featuring the more superficial habitats in which pulse-type fish live, wave-type EODs may be a better choice in the relatively simpler river-bed environments in which many wave-type fish live (e.g., the benthic zone of deep river channels; Crampton, 2019). In this case, achieving a good spatial resolution is critical during social encounters, especially considering the feeble nature of visual cues in low light conditions. In such habitats, social encounters may be less “abrupt”, but spatially less “conspicuous” or blurred.

Revised version: If this idea finds further support, the question arises as to why not all gymnotiforms electrolocate using the same strategy. It may be that a trade-off exists between space- and time-resolution in the evolution and maintenance of electrosensory systems (Crampton, 2019). While broadband pulse signals may be useful to capture highly complex environments rich in foliage, roots and other structures common in the more superficial habitats in which pulse-type fish live, wave-type EODs may be a better choice in the relatively simpler river-bed environments in which many wave-type fish live (e.g., the benthic zone of deep river channels; Crampton, 2019). In this case, achieving a good spatial resolution is critical during social encounters, especially considering the limited utility of visual cues in these low light conditions. In such habitats, social encounters may be less “abrupt”, but spatially less “conspicuous” or blurred (as a 3D electric field may be). In such a scenario, chirps could serve as a means to supplement the spatial information acquired via the beat, accentuating these cues during periods of reduced resolution.

Conclusions

Earlier version: While our results do not completely dismiss the possibility of chirps serving a role in social communication, the high stereotypy and auto-correlation of chirp time-series, along with the absence of any meaningful social correlate of chirping, suggests this scenario to be unlikely. Given the occurrence of chirping at a threshold distance and its increase in cluttered environments, this behavior mirrors the emission of other self-referenced signals used to gather more spatial information obtained from a nearby source (such as echolocation calls; Siemers et al., 2009). Could chirps be useful for both electrocommunication and electrolocation functions? Could chirps be imbued with different meanings?

Revised version: While our results do not completely dismiss the possibility of chirps serving a role in social communication, the high stereotypy and auto-correlation of chirp time-series, along with the absence of any meaningful social correlate of chirping, suggests this scenario to be unlikely. Given the clear distance threshold at which chirps are produced, the correlation with locomotion and the increase of chirping in cluttered environments, chirp production mirrors the emission of other self-referenced signals used to gather spatial information obtained from a nearby source (such as echolocation calls; Siemers et al., 2009). Could chirps be useful for both electrocommunication and electrolocation functions? Or more simply, could chirping be ascribed simultaneously to different functional tasks?

Signals with such dual function (i.e. echolocation and behavioral referencing) may be a rare find in nature (eavesdropped signals can become cues for localization purposes but are not intentionally produced with that aim), whereas it is more likely that signals directly affecting both senders and receivers do so by delivering the same type of information (i.e. spatial, in this case). In this scenario, probing cues could function simultaneously as proximity signals to signal presence, deter approaches, or coordinate behaviors like spawning, if properly timed (Henninger et al., 2018).

Conversely, if chirps would be referenced to specific behavioral states - as posed by the communication-hypothesis - the variable rate of chirp production across individuals and the infrequent production of chirp types other than type 2 chirps would render such a communication channel unreliable, semantically limited, and potentially incompatible with previously attributed functions of chirps.

In summary, although our behavioral experiments may not offer a comprehensive causal understanding of chirps and behavior, they do shed light on the potential and, in part, hitherto unexplored roles of chirping. These findings aim to stir the pot and initiate a discussion on possible alternative functions of chirps beyond their presumed communication role.

In echolocating species, the effectiveness of active sensing has been demonstrated to rely significantly on environmental complexity and spatial features (Fenton and Bell, 1981; Siemers et al., 2009; Fouda et al., 2018). Considering chirps as probing cues, knifefish might adaptively use different types of chirps based on the features of the foreign EOD source (frequency, position, location, orientation, distance, etc.). We emphasize the role of chirps as homeoactive signals - as opposed to alloactive - to underline the fact that chirps represent active modulations of signal frequency possibly aimed at optimizing the otherwise passive responses to the beat, occurring continuously and unintentionally (Zweifel and Hartmann, 2020). By offering a novel and more parsimonious explanation for chirping behavior, this study reignites the debate on the evolutionary reasons behind the development of this signaling system in these nocturnal bony fish, questioning whether it primarily serves electrolocation or communication purposes.

Materials and Methods

Animals

A total of 234 Apteronotus leptorhynchus of both sexes - age ranging between 2-3 years - were obtained from tropical fish importers and housed in individual 80 x 35 x 40 cm aquariums under a 12/12 light cycle. Fish were allocated to the different experiments as follows: 130 for the assessment of context-dependent effects on chirping, 16 fish were used for the playback experiments, 14 fish were used for the playback experiments involving EOD frequency ramps, 24 fish were used for the chirp-behavior correlative analysis, 30 fish were used in the novel environment exploration assay, 12 fish were used to assess the environment effects on chirping and 8 to assess the impact of clutter on the fish responses to playback EODs. In all groups both sexes were equally represented. Throughout the experiments, the water was continuously filtered, water conductivity was maintained between 150-300 µS cm-1, temperature between 23-26 °C and pH between 6-8. The hometank environment consisted of PVC shelter tubes and plastic plants. The fish were fed 3 times a week with red mosquito larvae. Prior to each experiment, the fish’s EOD frequency (EODf), body length and weight were measured. To identify their sex, the EODf was normalized to a water temperature of 25.0 °C using a Q10 of 2 (Dunlap and Ragazzi, 2015; Dunlap et al., 2000). The limitation of this approach is that females cannot be distinguished from immature males with absolute certainty, since no post-mortem gonadal inspection was carried out. Nevertheless, fish with a normalized EODf higher than 750 Hz are considered males and those with lower frequencies are considered females (Dunlap et al., 2000). Although a more accurate way to determine the sex of brown ghosts would be to consider other morphological features such as the shape of the snout (Hagedorn, 1986), the body size, the occurrence of developing eggs, EOD frequency has been extensively used for this purpose (Hagedorn, 1986; Meyer et al.,1987; Bastian et al., 2001).

Recordings of fish pairs and playback experiments

Equipment

All technical equipment used in the playback experiments is listed in the table below.

Experimental setup and groups

Fish pairs were recorded in 2 adjacent chambers of tripartite 80 x 35 x 40 cm aquariums (Figures 1, 2, 5). Tank compartments were separated by a plastic mesh barrier to allow electrical but no physical interaction. EOD recordings were conducted using 2 graphite electrodes placed on opposite sides of the tank (1 pair per compartment). EOD recordings were assigned to the following categories based on the type of experience fish had with either the test aquarium or the paired subjects: resident, intruder, dominant, subordinate, novel, experienced, courtship, no courtship, divided and free. Resident fish were housed for 1 week alone in the same setup, before being paired with intruder fish. Dominance was assessed by means of shelter-competition tests prior to EOD recordings. These tests consisted of 30 min long trials in which fish pairs were allowed to interact freely and compete for the occupancy of a plastic shelter tube positioned in one of the 3 compartments of a tripartite aquarium (1 pair per aquarium). The fish spending more time in the shelter tube was considered to be the dominant one, the other the subordinate. Fish were selected such that one fish was always 3-5 grams heavier than the other, to ensure a predictable outcome of the competition and to limit aggressive displays to a minimum (Jennions and Backwell 1996; Umbers et al. 2012). In these and other recordings, fish were considered “experienced” after at least 1 week of pairing, as opposed to “novel” (just paired). This means that residents, intruders, dominant and subordinate fish were all assigned to the category “novel”, whereas the category “experienced” included only resident and intruder fish, due to their longer pairing period (1 week). In a separate set of recordings, male and female pairs were subject to water conductivity changes to simulate the onset of the breeding season: during the course of a 4-week period water conductivity was lowered from 400 to 100 µS cm-1 (novel fish in high conductivity water = no breeding, experienced fish in low conductivity water = breeding season). Although, the term “breeding” here refers to the context and not to the actual behavioral repertoire displayed by the fish, this treatment resulted in 6/8 female fish to show evident signs of developed eggs (no post-mortem exam was used to confirm egg presence in the other 2 fish) while in other fish cohorts (housed in larger aquariums and with improved shelters) we successfully obtained newly spawned fish. At the end of this period, fish were allowed to swim freely and interact without mesh barriers (experienced fish in low conductivity water = breeding + free). Recordings from such pairs were compared with naïve male-female pairs freely swimming in high conductivity (400 µS cm-1) water (naïve fish in high conductivity water = no breeding + free). Although even in this case fish were paired with their tank partners for longer times, due to the different treatment they had been subject to, they were not included in the “experienced” category.

Chirp type categorization

Overall, the coding of custom MATLAB scripts (for EOD recording, chirp detection and validation), the preparation and execution of the behavioral experiments and the manual analysis of chirps imposed very significant timing constraints for the execution of the experiments. For this reason, to come up with threshold values usable to categorize chirps over the course of the whole study, we have used a preliminary dataset of 11342 chirps obtained from our first set of recordings (8 male-male fish pairs). A more complete dataset (N = 30486) was obtained and validated only 2 years later.

The cut-off values used in our study (50 ms duration and 105 Hz frequency modulation amplitude) were chosen based on the distribution of our preliminary recorded chirps but also based on reference values previously published by other authors (see Table 1, for previously published chirp categories and Table 3 for details on our own categorization).

All chirps recorded (N = 70219) were obtained from different experiments in these proportions: 30486 from staged fish pair interactions, 15720 from playback experiments, 3966 from ramp playback experiments, 4261 from freely swimming fish pairs used for chirp-behavior correlation analysis (297 of these were used for electric image calculations), 7893 from novel environment exploration experiments, 2051 from fish interactions in cluttered environments, 5842 from clutter playback experiments.

Amplifier and recording settings

EOD recordings (3 channels) were amplified through DPA-2FSL and DPA-2FS amplifiers (npi electronics) with a gain of 200, low pass filtered at 10 kHz, high-pass filtered at 100 Hz. The acquisition sampling rate was kept at 20 kHz in all recordings.

Playback experiments

All playback experiments were conducted in 80 x 35 x 40 cm glass aquariums in a dark room (Figure 3, 4). Water conductivity and temperature were similar to the fish housing conditions and ranged between 200-300 µS cm-1 and 23-26 °C, respectively. Water was pre-heated using commercially available heaters (EHEIM 3612 heating rod 50W). At one end of the tank, a mesh barrier separated the playback electrode from the fish. Electrodes were placed 10 cm apart and 1 cm away from the barrier. This layout ensured that playback stimuli consisted of naturalistic electric fields (Kelly et al., 2008). Electrodes were randomly placed at either side of the aquarium, in different experiments, to avoid playback location biases.

Three pairs of recording electrodes were placed in the fish compartment in accordance with an “IX” layout: 1 pair was parallel to the short tank side (“I”) the other 2 pairs were oriented diagonally across the compartment (“X”). This layout ensured that the fish EOD was always detectable, regardless of fish movement and body orientation. Prior to commencing the study, the stimulus amplitude was calibrated to have a field intensity of approximately 1 mV/cm measured across the mesh divider at circa 5 cm distance from the playback dipole, and in line with its axis.

Playback sinewave stimuli were designed in MATLAB and their frequency based on the measured fish EOD frequency prior to trial onset. Stimuli were then delivered via a 6211-USB DAQ (National Instruments) to the aquarium through a stimulus isolator (ISO-STIM 01M, npi electronics). The stimuli played back were modified sinewaves mimicking conspecific EODs the frequencies of which were calculated as differences (DF) from the fish’s own EOD frequency: -240, -160, -80, -40, -20, -10, -5, 0, 5, 10, 20, 40, 80, 160, 240 (Hz). Each of these playback stimuli was delivered in 4 different modes distinguished by their chirp content (modes, see Table 2 below): mode0 contained no chirps, mode1 contained type 1 chirps, mode2 with type 2 chirps and mode3 contained small rises (abrupt frequency rises). Each playback trial (50 s with a 5 s time of fade-in and fade-out) had a 180 s inter-trial interval. Each playback session consisted in a randomized sequence of 15 x 4 playback trials (1 min + 3 min each) which lasted for 4 h.

Video tracking

During playback experiments (Figure 3) but also freely swimming interactions (Figure 6) and novel environmental exploration assays (Figure 8), freely swimming fish were recorded at 40 FPS using an infrared USB camera (Grasshopper3, FLIR, model: GS3-U3-41C6NIR-C). The camera was set to acquire a frame every 500 samples of the EOD recording (1 sec = 20000 samples) through a TTL channel of the digital acquisition device (NI USB 6211 National Instruments). Swimming trajectories were extracted from the videos using the software BioTracker after converting the files from H264 to AVI (https://www.igb-berlin.de/en/biotracker). Playback experiments began around 1 p.m., after the onset of the animals lights OFF phase (12 AM lights ON, 12 PM lights OFF). Test subjects were acclimated for 30 min to the test aquarium, before trial onset.

Ethogram used to analyze freely swimming behavioral interactions

During freely swimming interactions, behaviors were annotated following previously described reports (Triefenbach and Zakon, 2008). A total of twelve fish pairs were first habituated for 30 minutes to the testing aquarium (30 x 80 cm) in presence of a mesh divider. Right after, the divider was removed and fish were allowed to interact freely for another 30 minutes. Both their EODs and behaviors were recorded in both sessions. During the habituation session, most fish remained in proximity to the barrier, parallel to it, swimming back and forth or stationary. Some fish remained stationary in a corner of their tank compartment, others swam more often exploring their surroundings. In the freely swimming session, fish often continued their exploration rounds before interacting more closely. This often implied a repeated circling of the whole tank contour. Interactions often started with backward swimming (a behavior we referred to as “knife”), often resolving in a reciprocal backward tail-to-tail approach (“tail1”). In some cases, one of the two fish would flicker the tail in rapid movements (“tail waving”, or “tail2”). From this positioning, fish would often slide onto each other to begin a more direct physical interaction (“wrestle”). This would often imply locomotion (“loc”) towards one end of the tank. When reaching a corner fish would either disengage and part ways or start an aggressive interaction, visible through either head butts (one fish hitting with the snout the flank of the other, ”head”) or - somehow more frequently - though bites directed to the head of the other fish (“atk”). Less commonly, the attacking fish would quiver their jaws (“jaw”) prior to either an attack (directed to the head or the flank) or a chasing attempt (“ch”). The number and extent of such encounters varied considerably across fish pairs: in some cases, it lasted for the whole duration of the recording, while in others fish were just resting at opposite sides of the tank.

Boundary element model and electric images

Added in the revised version: The results shown in Figure 7 (electric images) and Figure S12 (electric field) were calculated using the boundary element method (BEM; Pedraja et al., 2014; Hofmann et al., 2017; Rother et al., 2003). This model comprises two main components: a geometric reconstruction of the fish’s body and a calculation of the transcutaneous field by solving the Poisson equation for the fish boundary.

Added in the revised version: Briefly, the BEM method determines the boundary electrical distributions by solving a linear system of M x N equations, where M represents poles and N represents nodes. The unknown variables in this system are the trans-epithelial current density and potential at each node. These variables are computed for each node and linearly interpolated for the triangles defined by the nodes, which form the geometry of the fish and objects. Fish geometry as well as internal and skin conductivities were modeled using realistic information.

Added in the revised version: Furthermore, the BEM method allows for the concatenation of multiple instants, enabling the rendering of an entire sequence of electric fields and resulting electric images incorporating realistic electric organ discharge (EOD) characteristics such as waveform, duration, and frequencies (see Pedraja et al., 2014).

Added in the revised version: Here, electric images were computed for each fish position and chirp scenario by simulating various phases (the simulations started with four equally spaced phase relationships between the sender and receiver EODs, and were then averaged) and were calculated as the sum of the changes in current flow induced at the peaks of the amplitude-modulated EODs throughout the entire duration of the chirp (analysis window = 500 msec), maintaining consistent timing for both chirp and non-chirp conditions, across approximately 800 body nodes. The ΔEI images of Fig. 7C (chirps vs. beat) thus show the cumulative change in current flow over the entire body surface during a chirp compared to the same time window without a chirp.

Added in the revised version: Chirps utilized in the analysis were obtained from recordings of real fish pairs engaged in freely moving interactions (N = 4 fish pairs), with both fish producing chirps. The dataset consisted of 297 chirps, categorized into 87 type 1, 185 type 2, 27 type 3, and 4 rises.

Removed in the revised version: The results shown in Figure 7 (electric images) and Figure S12 (electric field) were calculated using the boundary element method (BEM; Pedraja et al., 2014; Hofmann et al., 2017; Rother et al., 2003); the electric images triggered by chirps were calculated as the sum of the absolute peak to peak amplitude of the beat at each body node (ca. 800 nodes). The EIs generated by interacting fish during chirps were compared with a basal condition (sender fish alone). Chirps used in this analysis were obtained from recordings of real fish pairs during freely moving interactions (N = 4 fish pairs) and consisted in a sample of 297 chirps (87 type 1, 185 type 2, 27 type 3 and 4 rises).

Beat decay estimates

The beat dissipation range in our experimental conditions was estimated by recording the interaction of two EOD mimics (a static reference electrode pair and another identical electrode pair, placed parallel to it, as a moving source) at increasing distances and in the same water conditions as in our recordings (conductivity 200 μS/cm and temperature 25°C; Figure S12). Both the reference and the moving EOD mimic were scaled to natural fish signals (1mV/cm, measured at 2-3 cm from the playback dipoles). Recordings were made placing recording electrodes at opposite poles of the reference electrodes while shifting the moving source with 5 cm steps up to 60 cm away. The EOD field intensity generated by both reference and moving electrodes was recorded for 5 seconds at each distance step, always at the extremities of the reference pair.

Environmental manipulations

Fish pairs were exposed to changing environmental conditions in the experiments illustrated in Figure 9. Recordings in lights ON, OFF and cluttered conditions were carried out in 80×35×40 cm aquariums divided in 2 compartments by a plastic mesh divider. During lights-ON, even illumination of both tank compartments was provided via a Leica LED illuminator located on a shelf above the aquarium (Leica CLS 150 XE Microscope Cold Light Source, 150W). During lights-OFF and cluttered conditions, a drape obtained from a blackout curtain was used to fully cover the tank (room lights were switched OFF and daylight lowered by curtains). Cluttered conditions were exactly the same as lights-OFF with the addition of floating and submerged plastic shelter tubes and plants heterogeneously arranged in order to fill up each compartment.

Data analysis

Chirp detection

Chirps were detected in a two-step process consisting of a first automatic detection followed by a manual validation. The first step implied a measurement of the power density in the FFT of the recorded signal (FFT parameters: window size = 2^12; overlap = 90%) within the range set by the fundamental frequency and the first harmonic component (EODf - 5 Hz and 2 x EOD - 100 Hz). Chirps were detected as peaks in the power density. The manual validation consisted of a survey of the detected chirps to eliminate false positives. False negatives were estimated at a rate lower than 10%. Although the two electrode pairs used in each aquarium would detect EODs from both fish, chirps were properly assigned to the sender based on signal intensity and the baseline EOD frequency of the two fish, when possible (EODs generated by a fish will have higher intensity if recorded by electrodes placed in the same compartment). Notably, other authors reported the occurrence of more than 2 types of large chirps (namely: type3-6), we included all those in one single group based on the distribution of our data and considering the absence of any obvious clustering for large chirps (see Figure 1).

Factor analysis mixed data - I

Factor analysis was conducted using the R package FactoMiner (“Analyse factorielle de donnees mixtes.” Revue Statistique Appliquee. LII (4). pp. 93-111. Josse and Husson, 2008).

Factor analysis of mixed data (FAMD, Figure 2) is equivalent to a principal component analysis adapted to analyze mixed datasets, containing both quantitative and qualitative variables (Pagès 2004) and explore their associations. The FAMD algorithm can be seen as a hybrid of principal component analysis (PCA) and multiple correspondence analysis (MCA): it works like PCA with quantitative variables and MCA for qualitative ones. Quantitative and qualitative variables are normalized in order to balance the influence of each set.

A total of 30486 sampled chirps were used. For each chirp, the following qualitative and quantitative variables were considered: frequency_modulation (peak of EOD instant frequency during chirp), chirp duration (instant frequency peak duration), EOD frequency of the sender (EODs), EOD frequency of the receiver (EODr), EOD_amplitude (amplitude of the EOD sinewave during chirp), power (chirp triggered EOD intensity), water temperature, water conductivity, chirp timestamp, weight and length of the sender fish, weight and length of the receiver fish, sender and receiver sex, status of the sender and receiver fish (experience with the recording aquarium), fish ID (to account for interindividual variability), chirp type (categorized as above), state (based on the experience with the other fish of the pair).

A first model (3 dimensions, model#1, dim1 17.46%, dim2 15.76%, dim3 10.67%) was made considering all variables listed above. Subsequent models were calculated while progressively reducing the number of variables (down to 3, model#11, dim1 23.17%, dim2 20.4%, dim3 14.35%). At each iteration variables with eigenvalues lower than the mean average contribution were eliminated.

Factor analysis mixed data - II

To assess the correlations among swimming related variables and chirps, a total of 7893 chirps obtained from a separate experiment (novel environment exploration, Figure 8) were used. For each chirp, the following qualitative and quantitative variables were considered: the percentage of time spent near the tank wall (“wall”), the percentage of time spent in the tank open space (“open”), percentage of time spent in the shelter area (“shelter”), the percentage of time spent within a rectangular area 10 cm around the conspecific location (“social”), the percentage of time spent investigating the novel object (“object”), the time at which chirps were produced (“timestamp”), the chirp type, the fish ID, the EOD frequency of the freely swimming fish, the EOD frequency of the caged fish, the sex of the chirp sender, the sex of the chirp receiver, the DF (i.e. EOD frequency difference between sender and receiver), the distance between sender and receiver (“distance”), the average speed of the sender (“avg speed”), the angle between sender and receiver (referred to the a horizontal axis delineated by the tube, “angle”), the interference estimate weighed on the distance (i.e. calculated on a beat resulting from EODs attenuated as a function of distance, “actual interference”) and the maximum beat interference possible for a given chirp (i.e. without distance attenuation, “interference”, Figure S11).

Chirp transitions and time-series cross-correlation

Earlier version: The number of chirp transitions present in each recording (dataset used for Figures 1, 2, 5) was measured by searching in a string array containing the 4 chirp types, all their possible pairwise permutations. The number of each chirp-type transition was then used to create a 4×4 frequency matrix. Chirp transition diagrams were obtained by calculating the median of such frequency matrices calculated for different fish pairs. Chirp time series were extracted from each recording and binned (50 ms) to evaluate the cross-correlation index within a +-2 s lag window.

Revised version: The number of chirp transitions present in each recording (dataset used for Figures 1, 2, 5) was measured by searching in a string array containing the 4 chirp types per fish pair, all their possible pairwise permutations (i.e. all possible permutations of 4+4=8 elements are: 1-1, 1-2, 1-3 … 7-6, 7-7, 7-8; considering the following legend 1 = fish1 type 1, 2 = fish 1 type 2, 3 = fish1 type 3 … 6 = fish2 type 2, 7 = fish2 type 3 and 8 = fish2 rise). The number of each chirp-type transition was then used to create a 4×4 frequency matrix. Chirp transition diagrams were obtained by calculating the median of such frequency matrices calculated for different fish pairs. Chirp time series were extracted from each recording and binned (50 ms) to evaluate the cross-correlation index within a +-2 s lag window.

Chirp interference

Chirp interference was calculated as the cumulative duration of outlier beat inter-peak intervals (IPI) after calculating the beat envelope for each chirp/DF combination (Figure S8). Gaussian chirps were generated on top of a constant frequency carrier of 862 Hz (arbitrarily chosen) and the beat resulting from a paired sinewave with a frequency within ±300 Hz of 862 Hz. The frequency modulation and duration of gaussian chirps were determined by setting the sigma and the amplitude of the peak and by measuring the actual peak size of the instant frequency (“medfreq” MATLAB function). This was obtained from the FFT analysis of the chirping signal (window size 2^12, 90% overlap, sampling frequency 20kHz). For each chirp, beat peaks (i.e. the peaks of the amplitude modulation, AM) were detected searching for local maxima in the signal first order envelope. Adjacent beat peaks were used to calculate beat cycle durations for the whole segment. The different impact of different types of chirps on the beat would result in different populations of beat cycles (depending on the beat frequency, the duration and FM of the chirp). Outliers in the beat cycle durations resulting from each chirp were used as a measure of how a given chirp would perturbate the otherwise regular beat cycles. Beat cycle outliers were detected using a MATLAB built-in function - “isoutlier” - with the “quartile” setting to include peak durations on both tails of the peak distribution (below 25% and above the 75% quartiles). The interference value for each chirp was obtained by dividing the cumulative outlier beat cycle duration (i.e. the total duration of beat cycles significantly affected by the chirp) by the overall duration of the beat peaks. The duration of each beat cycle corresponded to the difference in time between consecutive beat peaks. In this way, the effect of a chirp could be weighted on the number of peaks (i.e. the beat frequency). The beat interference is calculated within a fixed time window (700 ms, but see Figure S9 for wider windows) which corresponds to the median inter-chirp interval (ICI; median = 0.6987) calculated for 30486 chirps (130 chirp time-series) and is centered around the chirp (chirp peak is always at 350 msec). This parameter is used to provide an estimate of the effect of a chirp in its immediate temporal surround.

Text-mining analysis of chirp-behavior sequences

A text-mining approach was used to identify correlations between behaviors and chirp types in text strings containing both (pooled data from 12 fish pairs of mixed sexes) using KH Coder (Higuchi and Koichi, 2017). KH Coder is open-source software designed for computer-assisted qualitative data analysis, with a focus on quantitative content analysis, text mining, and applications in computational linguistics. After creating a word dataset including all the occurrences of chirps and behaviors in string form (words) conserving their temporal order, we have used KH Coder to perform a co-occurrence analysis and a multidimensional scaling. The first aims at analyzing text identifying potential relationships between words. By definition, co-occurrence networks are represented as interconnected networks based on the paired presence of words within a specified unit of text. The second, multidimensional scaling (MDS), is a means of visualizing the level of similarity of individual words based on their appearance pattern within a dataset. MDS is used to translate “information about the pairwise ‘distances’ among a set of words into a configuration of points mapped into an abstract cartesian space.

Statistics on chirp counts

The numbers of chirps produced during the experiments assessing the effect of environmental factors on chirp rate production were compared using the Friedman’s test (non-parametric ANOVA for repeated measures) as data were not following a Gaussian distribution.

Added in the revised version: Supporting information